9
15816
210
UNSAT
tff(type_def_5, type, 'Nat_1()': $tType).
tff(type_def_6, type, 'Bool_0()': $tType).
tff(type_def_7, type, 'list_0()': $tType).
tff(type_def_8, type, 'Nat_0()': $tType).
tff(func_def_0, type, 'Z_12': 'Nat_1()').
tff(func_def_1, type, 'Z_13': ('Nat_1()') > 'Nat_1()').
tff(func_def_2, type, false_0: 'Bool_0()').
tff(func_def_3, type, true_0: 'Bool_0()').
tff(func_def_4, type, nil_0: 'list_0()').
tff(func_def_5, type, cons_0: ('Nat_1()' * 'list_0()') > 'list_0()').
tff(func_def_6, type, zero_0: 'Nat_0()').
tff(func_def_7, type, succ_0: ('Nat_0()') > 'Nat_0()').
tff(pred_def_1, type, diseqNat_1: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_2, type, unS_1: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_3, type, isZ_2: ('Nat_1()') > $o).
tff(pred_def_4, type, isZ_3: ('Nat_1()') > $o).
tff(pred_def_5, type, add_0: ('Nat_1()' * 'Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_6, type, minus_1: ('Nat_1()' * 'Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_7, type, le_0: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_8, type, ge_0: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_9, type, lt_1: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_10, type, gt_0: ('Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_11, type, mult_0: ('Nat_1()' * 'Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_12, type, div_0: ('Nat_1()' * 'Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_13, type, mod_1: ('Nat_1()' * 'Nat_1()' * 'Nat_1()') > $o).
tff(pred_def_14, type, diseqBool_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_15, type, isfalse_1: ('Bool_0()') > $o).
tff(pred_def_16, type, istrue_1: ('Bool_0()') > $o).
tff(pred_def_17, type, and_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_18, type, or_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_19, type, hence_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_20, type, not_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_21, type, diseqlist_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_22, type, head_1: ('Nat_1()' * 'list_0()') > $o).
tff(pred_def_23, type, tail_1: ('list_0()' * 'list_0()') > $o).
tff(pred_def_24, type, isnil_0: ('list_0()') > $o).
tff(pred_def_25, type, iscons_0: ('list_0()') > $o).
tff(pred_def_26, type, diseqNat_0: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_27, type, p_1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_28, type, iszero_0: ('Nat_0()') > $o).
tff(pred_def_29, type, issucc_0: ('Nat_0()') > $o).
tff(pred_def_30, type, take_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_31, type, plus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_32, type, minus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_33, type, lt_0: ('Bool_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_34, type, mod_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_35, type, length_0: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_36, type, drop_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_37, type, x_7: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_38, type, rotate_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_39, type, sP0: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_40, type, sP1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_41, type, sP2: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_42, type, sP3: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_43, type, sP4: ('Nat_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_44, type, sP5: ('Nat_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_45, type, sP6: ('Nat_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_46, type, sP7: ('Nat_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_47, type, sP8: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(f8754,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f722,f4715,f5095,f5097,f258])).
tff(f258,plain,(
  ( ! [X4 : 'list_0()',X2 : 'Nat_0()',X7 : 'list_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~sP8(X4,X2,X7) | ~sP5(X2,X7,X1) | ~sP6(X2,X7,X9) | ~x_7(X9,X1,X4)) )),
  inference(general_splitting,[],[f256,f257_D])).
tff(f257,plain,(
  ( ! [X4 : 'list_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()'] : (~sP7(X8,X2,X7) | ~take_0(X4,X8,X7) | sP8(X4,X2,X7)) )),
  inference(cnf_transformation,[],[f257_D])).
tff(f257_D,plain,(
  ( ! [X7,X2,X4] : (( ! [X8] : (~sP7(X8,X2,X7) | ~take_0(X4,X8,X7)) ) <=> ~sP8(X4,X2,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP8])])).
tff(f256,plain,(
  ( ! [X4 : 'list_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~take_0(X4,X8,X7) | ~x_7(X9,X1,X4) | ~sP5(X2,X7,X1) | ~sP6(X2,X7,X9) | ~sP7(X8,X2,X7)) )),
  inference(general_splitting,[],[f254,f255_D])).
tff(f255,plain,(
  ( ! [X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X3 : 'Nat_0()'] : (~mod_0(X8,X2,X3) | ~length_0(X3,X7) | sP7(X8,X2,X7)) )),
  inference(cnf_transformation,[],[f255_D])).
tff(f255_D,plain,(
  ( ! [X7,X2,X8] : (( ! [X3] : (~mod_0(X8,X2,X3) | ~length_0(X3,X7)) ) <=> ~sP7(X8,X2,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP7])])).
tff(f254,plain,(
  ( ! [X4 : 'list_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~length_0(X3,X7) | ~mod_0(X8,X2,X3) | ~take_0(X4,X8,X7) | ~x_7(X9,X1,X4) | ~sP5(X2,X7,X1) | ~sP6(X2,X7,X9)) )),
  inference(general_splitting,[],[f252,f253_D])).
tff(f253,plain,(
  ( ! [X2 : 'Nat_0()',X7 : 'list_0()',X5 : 'list_0()',X9 : 'list_0()'] : (~rotate_0(X5,X2,X7) | ~diseqlist_0(X5,X9) | sP6(X2,X7,X9)) )),
  inference(cnf_transformation,[],[f253_D])).
tff(f253_D,plain,(
  ( ! [X9,X7,X2] : (( ! [X5] : (~rotate_0(X5,X2,X7) | ~diseqlist_0(X5,X9)) ) <=> ~sP6(X2,X7,X9)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP6])])).
tff(f252,plain,(
  ( ! [X4 : 'list_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~diseqlist_0(X5,X9) | ~rotate_0(X5,X2,X7) | ~length_0(X3,X7) | ~mod_0(X8,X2,X3) | ~take_0(X4,X8,X7) | ~x_7(X9,X1,X4) | ~sP5(X2,X7,X1)) )),
  inference(general_splitting,[],[f250,f251_D])).
tff(f251,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X7 : 'list_0()',X1 : 'list_0()'] : (~sP4(X6,X2,X7) | ~drop_0(X1,X6,X7) | sP5(X2,X7,X1)) )),
  inference(cnf_transformation,[],[f251_D])).
tff(f251_D,plain,(
  ( ! [X1,X7,X2] : (( ! [X6] : (~sP4(X6,X2,X7) | ~drop_0(X1,X6,X7)) ) <=> ~sP5(X2,X7,X1)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP5])])).
tff(f250,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'list_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~diseqlist_0(X5,X9) | ~rotate_0(X5,X2,X7) | ~drop_0(X1,X6,X7) | ~length_0(X3,X7) | ~mod_0(X8,X2,X3) | ~take_0(X4,X8,X7) | ~x_7(X9,X1,X4) | ~sP4(X6,X2,X7)) )),
  inference(general_splitting,[],[f240,f249_D])).
tff(f249,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'list_0()'] : (~mod_0(X6,X2,X0) | ~length_0(X0,X7) | sP4(X6,X2,X7)) )),
  inference(cnf_transformation,[],[f249_D])).
tff(f249_D,plain,(
  ( ! [X7,X2,X6] : (( ! [X0] : (~mod_0(X6,X2,X0) | ~length_0(X0,X7)) ) <=> ~sP4(X6,X2,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP4])])).
tff(f240,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'list_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()',X9 : 'list_0()'] : (~diseqlist_0(X5,X9) | ~rotate_0(X5,X2,X7) | ~length_0(X0,X7) | ~mod_0(X6,X2,X0) | ~drop_0(X1,X6,X7) | ~length_0(X3,X7) | ~mod_0(X8,X2,X3) | ~take_0(X4,X8,X7) | ~x_7(X9,X1,X4)) )),
  inference(cnf_transformation,[],[f155])).
tff(f155,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'list_0()',X8 : 'Nat_0()',X9 : 'list_0()'] : (~diseqlist_0(X5,X9) | ~rotate_0(X5,X2,X7) | ~length_0(X0,X7) | ~mod_0(X6,X2,X0) | ~drop_0(X1,X6,X7) | ~length_0(X3,X7) | ~mod_0(X8,X2,X3) | ~take_0(X4,X8,X7) | ~x_7(X9,X1,X4))),
  inference(ennf_transformation,[],[f119])).
tff(f119,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'list_0()',X8 : 'Nat_0()',X9 : 'list_0()'] : ~(diseqlist_0(X5,X9) & rotate_0(X5,X2,X7) & length_0(X0,X7) & mod_0(X6,X2,X0) & drop_0(X1,X6,X7) & length_0(X3,X7) & mod_0(X8,X2,X3) & take_0(X4,X8,X7) & x_7(X9,X1,X4))),
  inference(true_and_false_elimination,[],[f118])).
tff(f118,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'list_0()',X8 : 'Nat_0()',X9 : 'list_0()'] : ((diseqlist_0(X5,X9) & rotate_0(X5,X2,X7) & length_0(X0,X7) & mod_0(X6,X2,X0) & drop_0(X1,X6,X7) & length_0(X3,X7) & mod_0(X8,X2,X3) & take_0(X4,X8,X7) & x_7(X9,X1,X4)) => $false)),
  inference(rectify,[],[f85])).
tff(f85,axiom,(
  ! [X1 : 'Nat_0()',X3 : 'list_0()',X8 : 'Nat_0()',X4 : 'Nat_0()',X6 : 'list_0()',X0 : 'list_0()',X2 : 'Nat_0()',X9 : 'list_0()',X5 : 'Nat_0()',X7 : 'list_0()'] : ((diseqlist_0(X0,X7) & rotate_0(X0,X8,X9) & length_0(X1,X9) & mod_0(X2,X8,X1) & drop_0(X3,X2,X9) & length_0(X4,X9) & mod_0(X5,X8,X4) & take_0(X6,X5,X9) & x_7(X7,X3,X6)) => $false)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f5097,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'Nat_0()',X1 : 'Nat_1()'] : (sP5(succ_0(succ_0(succ_0(X0))),cons_0(X1,cons_0(X2,nil_0)),cons_0(X1,cons_0(X2,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f200,f4889,f251])).
tff(f4889,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'Nat_0()',X1 : 'Nat_1()'] : (sP4(zero_0,succ_0(succ_0(succ_0(X0))),cons_0(X1,cons_0(X2,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f519,f4740,f249])).
tff(f4740,plain,(
  ( ! [X0 : 'Nat_0()'] : (mod_0(zero_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f529,f348,f309,f244])).
tff(f244,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~mod_0(X1,X2,succ_0(X4)) | ~minus_0(X2,X0,succ_0(X4)) | mod_0(X1,X0,succ_0(X4)) | ~sP1(X4,X0)) )),
  inference(general_splitting,[],[f238,f243_D])).
tff(f243,plain,(
  ( ! [X4 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Bool_0()'] : (~lt_0(X3,X0,succ_0(X4)) | ~diseqBool_0(X3,true_0) | sP1(X4,X0)) )),
  inference(cnf_transformation,[],[f243_D])).
tff(f243_D,plain,(
  ( ! [X0,X4] : (( ! [X3] : (~lt_0(X3,X0,succ_0(X4)) | ~diseqBool_0(X3,true_0)) ) <=> ~sP1(X4,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP1])])).
tff(f238,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Bool_0()',X1 : 'Nat_0()'] : (mod_0(X1,X0,succ_0(X4)) | ~diseqBool_0(X3,true_0) | ~minus_0(X2,X0,succ_0(X4)) | ~mod_0(X1,X2,succ_0(X4)) | ~lt_0(X3,X0,succ_0(X4))) )),
  inference(cnf_transformation,[],[f153])).
tff(f153,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : (mod_0(X1,X0,succ_0(X4)) | ~diseqBool_0(X3,true_0) | ~minus_0(X2,X0,succ_0(X4)) | ~mod_0(X1,X2,succ_0(X4)) | ~lt_0(X3,X0,succ_0(X4)))),
  inference(flattening,[],[f152])).
tff(f152,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : (mod_0(X1,X0,succ_0(X4)) | (~diseqBool_0(X3,true_0) | ~minus_0(X2,X0,succ_0(X4)) | ~mod_0(X1,X2,succ_0(X4)) | ~lt_0(X3,X0,succ_0(X4))))),
  inference(ennf_transformation,[],[f115])).
tff(f115,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : ((diseqBool_0(X3,true_0) & minus_0(X2,X0,succ_0(X4)) & mod_0(X1,X2,succ_0(X4)) & lt_0(X3,X0,succ_0(X4))) => mod_0(X1,X0,succ_0(X4)))),
  inference(rectify,[],[f69])).
tff(f69,axiom,(
  ! [X4 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Bool_0()',X3 : 'Nat_0()'] : ((diseqBool_0(X2,true_0) & minus_0(X1,X4,succ_0(X3)) & mod_0(X0,X1,succ_0(X3)) & lt_0(X2,X4,succ_0(X3))) => mod_0(X0,X4,succ_0(X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f309,plain,(
  ( ! [X0 : 'Nat_0()'] : (mod_0(zero_0,zero_0,succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f205,f219])).
tff(f219,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~lt_0(true_0,X0,succ_0(X1)) | mod_0(X0,X0,succ_0(X1))) )),
  inference(cnf_transformation,[],[f128])).
tff(f128,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (mod_0(X0,X0,succ_0(X1)) | ~lt_0(true_0,X0,succ_0(X1)))),
  inference(ennf_transformation,[],[f68])).
tff(f68,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (lt_0(true_0,X0,succ_0(X1)) => mod_0(X0,X0,succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f205,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(true_0,zero_0,succ_0(X0))) )),
  inference(cnf_transformation,[],[f66])).
tff(f66,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(true_0,zero_0,succ_0(X0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f348,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f346,f225])).
tff(f225,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (minus_0(X1,succ_0(X2),succ_0(X0)) | ~minus_0(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f134])).
tff(f134,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X1,succ_0(X2),succ_0(X0)) | ~minus_0(X1,X2,X0))),
  inference(ennf_transformation,[],[f101])).
tff(f101,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X1,X2,X0) => minus_0(X1,succ_0(X2),succ_0(X0)))),
  inference(rectify,[],[f62])).
tff(f62,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X0,X2,X1) => minus_0(X0,succ_0(X2),succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f346,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(X0)),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f204,f225])).
tff(f204,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f63])).
tff(f63,axiom,(
  ! [X0 : 'Nat_0()'] : minus_0(zero_0,succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f529,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP1(succ_0(zero_0),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f163,f354,f243])).
tff(f354,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(succ_0(X0)),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f351,f226])).
tff(f226,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Bool_0()',X1 : 'Nat_0()'] : (lt_0(X0,succ_0(X2),succ_0(X1)) | ~lt_0(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f135])).
tff(f135,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (lt_0(X0,succ_0(X2),succ_0(X1)) | ~lt_0(X0,X2,X1))),
  inference(ennf_transformation,[],[f102])).
tff(f102,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (lt_0(X0,X2,X1) => lt_0(X0,succ_0(X2),succ_0(X1)))),
  inference(rectify,[],[f65])).
tff(f65,axiom,(
  ! [X0 : 'Bool_0()',X2 : 'Nat_0()',X1 : 'Nat_0()'] : (lt_0(X0,X1,X2) => lt_0(X0,succ_0(X1),succ_0(X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f351,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(X0),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f195,f226])).
tff(f195,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,X0,zero_0)) )),
  inference(cnf_transformation,[],[f67])).
tff(f67,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(false_0,X0,zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f163,plain,(
  diseqBool_0(false_0,true_0)),
  inference(cnf_transformation,[],[f27])).
tff(f27,axiom,(
  diseqBool_0(false_0,true_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f519,plain,(
  ( ! [X0 : 'Nat_1()',X1 : 'Nat_1()'] : (length_0(succ_0(succ_0(zero_0)),cons_0(X0,cons_0(X1,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f515,f241])).
tff(f241,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'list_0()',X3 : 'Nat_0()'] : (sP0(X3,X0) | length_0(X3,cons_0(X2,X0))) )),
  inference(cnf_transformation,[],[f241_D])).
tff(f241_D,plain,(
  ( ! [X0,X3] : (( ! [X2] : length_0(X3,cons_0(X2,X0)) ) <=> ~sP0(X3,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP0])])).
tff(f515,plain,(
  ( ! [X0 : 'Nat_1()'] : (~sP0(succ_0(succ_0(zero_0)),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f317,f513,f242])).
tff(f242,plain,(
  ( ! [X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~plus_0(X3,succ_0(zero_0),X1) | ~length_0(X1,X0) | ~sP0(X3,X0)) )),
  inference(general_splitting,[],[f236,f241_D])).
tff(f236,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(X3,cons_0(X2,X0)) | ~length_0(X1,X0) | ~plus_0(X3,succ_0(zero_0),X1)) )),
  inference(cnf_transformation,[],[f149])).
tff(f149,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_1()',X3 : 'Nat_0()'] : (length_0(X3,cons_0(X2,X0)) | ~length_0(X1,X0) | ~plus_0(X3,succ_0(zero_0),X1))),
  inference(flattening,[],[f148])).
tff(f148,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_1()',X3 : 'Nat_0()'] : (length_0(X3,cons_0(X2,X0)) | (~length_0(X1,X0) | ~plus_0(X3,succ_0(zero_0),X1)))),
  inference(ennf_transformation,[],[f113])).
tff(f113,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_1()',X3 : 'Nat_0()'] : ((length_0(X1,X0) & plus_0(X3,succ_0(zero_0),X1)) => length_0(X3,cons_0(X2,X0)))),
  inference(rectify,[],[f71])).
tff(f71,axiom,(
  ! [X3 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_1()',X0 : 'Nat_0()'] : ((length_0(X1,X3) & plus_0(X0,succ_0(zero_0),X1)) => length_0(X0,cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f513,plain,(
  ( ! [X0 : 'Nat_1()'] : (length_0(succ_0(zero_0),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f510,f241])).
tff(f510,plain,(
  ~sP0(succ_0(zero_0),nil_0)),
  inference(unit_resulting_resolution,[],[f165,f317,f242])).
tff(f165,plain,(
  length_0(zero_0,nil_0)),
  inference(cnf_transformation,[],[f72])).
tff(f72,axiom,(
  length_0(zero_0,nil_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f317,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f198,f224])).
tff(f224,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(X1),X2) | ~plus_0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f133])).
tff(f133,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(X1),X2) | ~plus_0(X0,X1,X2))),
  inference(ennf_transformation,[],[f60])).
tff(f60,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X0,X1,X2) => plus_0(succ_0(X0),succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f198,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f61])).
tff(f61,axiom,(
  ! [X0 : 'Nat_0()'] : plus_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f200,plain,(
  ( ! [X0 : 'list_0()'] : (drop_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f75])).
tff(f75,axiom,(
  ! [X0 : 'list_0()'] : drop_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f5095,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'Nat_0()',X1 : 'Nat_1()'] : (sP8(nil_0,succ_0(succ_0(succ_0(X0))),cons_0(X1,cons_0(X2,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f192,f4888,f257])).
tff(f4888,plain,(
  ( ! [X2 : 'Nat_1()',X0 : 'Nat_0()',X1 : 'Nat_1()'] : (sP7(zero_0,succ_0(succ_0(succ_0(X0))),cons_0(X1,cons_0(X2,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f519,f4740,f255])).
tff(f192,plain,(
  ( ! [X0 : 'list_0()'] : (take_0(nil_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f59])).
tff(f59,axiom,(
  ! [X0 : 'list_0()'] : take_0(nil_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f4715,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_1()',X1 : 'Nat_1()'] : (sP6(succ_0(succ_0(succ_0(zero_0))),cons_0(X0,cons_0('Z_12',nil_0)),cons_0('Z_13'(X1),X2))) )),
  inference(unit_resulting_resolution,[],[f362,f4481,f253])).
tff(f4481,plain,(
  ( ! [X0 : 'Nat_1()',X1 : 'Nat_1()'] : (rotate_0(cons_0(X0,cons_0(X1,nil_0)),succ_0(succ_0(succ_0(zero_0))),cons_0(X1,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f721,f4400,f237])).
tff(f237,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'Nat_1()',X3 : 'Nat_0()',X1 : 'list_0()'] : (~x_7(X2,X1,cons_0(X0,nil_0)) | rotate_0(X4,succ_0(X3),cons_0(X0,X1)) | ~rotate_0(X4,X3,X2)) )),
  inference(cnf_transformation,[],[f151])).
tff(f151,plain,(
  ! [X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()'] : (rotate_0(X4,succ_0(X3),cons_0(X0,X1)) | ~x_7(X2,X1,cons_0(X0,nil_0)) | ~rotate_0(X4,X3,X2))),
  inference(flattening,[],[f150])).
tff(f150,plain,(
  ! [X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()'] : (rotate_0(X4,succ_0(X3),cons_0(X0,X1)) | (~x_7(X2,X1,cons_0(X0,nil_0)) | ~rotate_0(X4,X3,X2)))),
  inference(ennf_transformation,[],[f114])).
tff(f114,plain,(
  ! [X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()'] : ((x_7(X2,X1,cons_0(X0,nil_0)) & rotate_0(X4,X3,X2)) => rotate_0(X4,succ_0(X3),cons_0(X0,X1)))),
  inference(rectify,[],[f78])).
tff(f78,axiom,(
  ! [X2 : 'Nat_1()',X3 : 'list_0()',X1 : 'list_0()',X4 : 'Nat_0()',X0 : 'list_0()'] : ((x_7(X1,X3,cons_0(X2,nil_0)) & rotate_0(X0,X4,X1)) => rotate_0(X0,succ_0(X4),cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f4400,plain,(
  ( ! [X0 : 'Nat_1()',X1 : 'Nat_1()'] : (rotate_0(cons_0(X0,cons_0(X1,nil_0)),succ_0(succ_0(zero_0)),cons_0(X0,cons_0(X1,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f721,f4324,f237])).
tff(f4324,plain,(
  ( ! [X0 : 'Nat_1()',X1 : 'Nat_1()'] : (rotate_0(cons_0(X0,cons_0(X1,nil_0)),succ_0(zero_0),cons_0(X1,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f199,f721,f237])).
tff(f199,plain,(
  ( ! [X0 : 'list_0()'] : (rotate_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f80])).
tff(f80,axiom,(
  ! [X0 : 'list_0()'] : rotate_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f721,plain,(
  ( ! [X0 : 'Nat_1()',X1 : 'list_0()'] : (x_7(cons_0(X0,X1),cons_0(X0,nil_0),X1)) )),
  inference(unit_resulting_resolution,[],[f197,f229])).
tff(f229,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_1()',X3 : 'list_0()',X1 : 'list_0()'] : (x_7(cons_0(X0,X2),cons_0(X0,X3),X1) | ~x_7(X2,X3,X1)) )),
  inference(cnf_transformation,[],[f138])).
tff(f138,plain,(
  ! [X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()'] : (x_7(cons_0(X0,X2),cons_0(X0,X3),X1) | ~x_7(X2,X3,X1))),
  inference(ennf_transformation,[],[f105])).
tff(f105,plain,(
  ! [X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()'] : (x_7(X2,X3,X1) => x_7(cons_0(X0,X2),cons_0(X0,X3),X1))),
  inference(rectify,[],[f76])).
tff(f76,axiom,(
  ! [X1 : 'Nat_1()',X3 : 'list_0()',X0 : 'list_0()',X2 : 'list_0()'] : (x_7(X0,X2,X3) => x_7(cons_0(X1,X0),cons_0(X1,X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f197,plain,(
  ( ! [X0 : 'list_0()'] : (x_7(X0,nil_0,X0)) )),
  inference(cnf_transformation,[],[f77])).
tff(f77,axiom,(
  ! [X0 : 'list_0()'] : x_7(X0,nil_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f362,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X1 : 'Nat_1()'] : (diseqlist_0(cons_0('Z_12',X0),cons_0('Z_13'(X1),X2))) )),
  inference(unit_resulting_resolution,[],[f184,f227])).
tff(f227,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_1()',X1 : 'Nat_1()'] : (diseqlist_0(cons_0(X1,X2),cons_0(X3,X0)) | ~diseqNat_1(X1,X3)) )),
  inference(cnf_transformation,[],[f136])).
tff(f136,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_1()',X2 : 'list_0()',X3 : 'Nat_1()'] : (diseqlist_0(cons_0(X1,X2),cons_0(X3,X0)) | ~diseqNat_1(X1,X3))),
  inference(ennf_transformation,[],[f103])).
tff(f103,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_1()',X2 : 'list_0()',X3 : 'Nat_1()'] : (diseqNat_1(X1,X3) => diseqlist_0(cons_0(X1,X2),cons_0(X3,X0)))),
  inference(rectify,[],[f49])).
tff(f49,axiom,(
  ! [X3 : 'list_0()',X0 : 'Nat_1()',X1 : 'list_0()',X2 : 'Nat_1()'] : (diseqNat_1(X0,X2) => diseqlist_0(cons_0(X0,X1),cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f184,plain,(
  ( ! [X0 : 'Nat_1()'] : (diseqNat_1('Z_12','Z_13'(X0))) )),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  ! [X0 : 'Nat_1()'] : diseqNat_1('Z_12','Z_13'(X0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/rotate_mod.smt2',unknown)).
tff(f722,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_1()',X1 : 'Nat_1()'] : (x_7(cons_0(X0,cons_0(X1,X2)),cons_0(X0,cons_0(X1,nil_0)),X2)) )),
  inference(unit_resulting_resolution,[],[f721,f229])).
