11
19900
460
UNSAT
tff(type_def_5, type, 'Bool_0()': $tType).
tff(type_def_6, type, 'Nat_0()': $tType).
tff(type_def_7, type, 'list_0()': $tType).
tff(type_def_8, type, 'pair_0()': $tType).
tff(func_def_0, type, false_0: 'Bool_0()').
tff(func_def_1, type, true_0: 'Bool_0()').
tff(func_def_2, type, zero_0: 'Nat_0()').
tff(func_def_3, type, succ_0: ('Nat_0()') > 'Nat_0()').
tff(func_def_4, type, nil_0: 'list_0()').
tff(func_def_5, type, cons_0: ('Nat_0()' * 'list_0()') > 'list_0()').
tff(func_def_6, type, pair_1: ('list_0()' * 'list_0()') > 'pair_0()').
tff(pred_def_1, type, diseqBool_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_2, type, isfalse_1: ('Bool_0()') > $o).
tff(pred_def_3, type, istrue_1: ('Bool_0()') > $o).
tff(pred_def_4, type, and_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_5, type, or_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_6, type, hence_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_7, type, not_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_8, type, diseqNat_0: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_9, type, p_1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_10, type, iszero_0: ('Nat_0()') > $o).
tff(pred_def_11, type, issucc_0: ('Nat_0()') > $o).
tff(pred_def_12, type, diseqlist_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_13, type, head_1: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_14, type, tail_1: ('list_0()' * 'list_0()') > $o).
tff(pred_def_15, type, isnil_0: ('list_0()') > $o).
tff(pred_def_16, type, iscons_0: ('list_0()') > $o).
tff(pred_def_17, type, diseqpair_0: ('pair_0()' * 'pair_0()') > $o).
tff(pred_def_18, type, projpair_2: ('list_0()' * 'pair_0()') > $o).
tff(pred_def_19, type, projpair_3: ('list_0()' * 'pair_0()') > $o).
tff(pred_def_20, type, ispair_0: ('pair_0()') > $o).
tff(pred_def_21, type, take_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_22, type, plus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_23, type, times_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_24, type, minus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_25, type, lt_0: ('Bool_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_26, type, leq_0: ('Bool_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_27, type, sort_0: ('list_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_28, type, length_0: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_29, type, insert_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_30, type, isort_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_31, type, idiv_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_32, type, drop_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_33, type, splitAt_0: ('pair_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_34, type, x_14: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_35, type, stoogesort_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_36, type, stoogesort_1: ('list_0()' * 'list_0()') > $o).
tff(pred_def_37, type, stoogesort_2: ('list_0()' * 'list_0()') > $o).
tff(pred_def_38, type, sP0: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_39, type, sP1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_40, type, sP2: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_41, type, sP3: ('list_0()' * 'list_0()') > $o).
tff(pred_def_42, type, sP4: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_43, type, sP5: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_44, type, sP6: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_45, type, sP7: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_46, type, sP8: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_47, type, sP9: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_48, type, sP10: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_49, type, sP11: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_50, type, sP12: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_51, type, sP13: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_52, type, sP14: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_53, type, sP15: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_54, type, sP16: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_55, type, sP17: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_56, type, sP18: ('list_0()' * 'Nat_0()') > $o).
tff(pred_def_57, type, sP19: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_58, type, sP20: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(f21992,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f7104,f4396,f21953,f266])).
tff(f266,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (~insert_0(X0,X3,X2) | ~isort_0(X2,X1) | isort_0(X0,cons_0(X3,X1))) )),
  inference(cnf_transformation,[],[f170])).
tff(f170,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (isort_0(X0,cons_0(X3,X1)) | ~isort_0(X2,X1) | ~insert_0(X0,X3,X2))),
  inference(flattening,[],[f169])).
tff(f169,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (isort_0(X0,cons_0(X3,X1)) | (~isort_0(X2,X1) | ~insert_0(X0,X3,X2)))),
  inference(ennf_transformation,[],[f124])).
tff(f124,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : ((isort_0(X2,X1) & insert_0(X0,X3,X2)) => isort_0(X0,cons_0(X3,X1)))),
  inference(rectify,[],[f61])).
tff(f61,axiom,(
  ! [X0 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()',X2 : 'Nat_0()'] : ((isort_0(X1,X3) & insert_0(X0,X2,X1)) => isort_0(X0,cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f21953,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X1 : 'Nat_0()'] : (~isort_0(cons_0(zero_0,X0),cons_0(succ_0(X1),cons_0(zero_0,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f427,f21947,f254])).
tff(f254,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X1 : 'list_0()'] : (~stoogesort_1(X2,X1) | ~diseqlist_0(X2,X0) | ~isort_0(X0,X1)) )),
  inference(cnf_transformation,[],[f155])).
tff(f155,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()'] : (~diseqlist_0(X2,X0) | ~stoogesort_1(X2,X1) | ~isort_0(X0,X1))),
  inference(ennf_transformation,[],[f111])).
tff(f111,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()'] : ~(diseqlist_0(X2,X0) & stoogesort_1(X2,X1) & isort_0(X0,X1))),
  inference(true_and_false_elimination,[],[f110])).
tff(f110,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()'] : ((diseqlist_0(X2,X0) & stoogesort_1(X2,X1) & isort_0(X0,X1)) => $false)),
  inference(rectify,[],[f89])).
tff(f89,axiom,(
  ! [X1 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()'] : ((diseqlist_0(X0,X1) & stoogesort_1(X0,X2) & isort_0(X1,X2)) => $false)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f21947,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_1(cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f13202,f21862,f285])).
tff(f285,plain,(
  ( ! [X6 : 'list_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~stoogesort_0(X3,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | stoogesort_1(X6,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~sP3(X3,X6)) )),
  inference(general_splitting,[],[f271,f284_D])).
tff(f284,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X3 : 'list_0()'] : (~stoogesort_2(X4,X3) | ~stoogesort_0(X6,X4) | sP3(X3,X6)) )),
  inference(cnf_transformation,[],[f284_D])).
tff(f284_D,plain,(
  ( ! [X6,X3] : (( ! [X4] : (~stoogesort_2(X4,X3) | ~stoogesort_0(X6,X4)) ) <=> ~sP3(X3,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP3])])).
tff(f271,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_1(X6,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~stoogesort_0(X3,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~stoogesort_2(X4,X3) | ~stoogesort_0(X6,X4)) )),
  inference(cnf_transformation,[],[f180])).
tff(f180,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_1(X6,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~stoogesort_0(X3,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~stoogesort_2(X4,X3) | ~stoogesort_0(X6,X4))),
  inference(flattening,[],[f179])).
tff(f179,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_1(X6,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | (~stoogesort_0(X3,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) | ~stoogesort_2(X4,X3) | ~stoogesort_0(X6,X4)))),
  inference(ennf_transformation,[],[f129])).
tff(f129,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : ((stoogesort_0(X3,cons_0(X5,cons_0(X0,cons_0(X2,X1)))) & stoogesort_2(X4,X3) & stoogesort_0(X6,X4)) => stoogesort_1(X6,cons_0(X5,cons_0(X0,cons_0(X2,X1)))))),
  inference(rectify,[],[f72])).
tff(f72,axiom,(
  ! [X5 : 'Nat_0()',X4 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X6 : 'Nat_0()',X0 : 'list_0()'] : ((stoogesort_0(X1,cons_0(X6,cons_0(X5,cons_0(X3,X4)))) & stoogesort_2(X2,X1) & stoogesort_0(X0,X2)) => stoogesort_1(X0,cons_0(X6,cons_0(X5,cons_0(X3,X4)))))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f21862,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP3(cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f13202,f21847,f284])).
tff(f21847,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_2(cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f852,f7742,f8244])).
tff(f8244,plain,(
  ( ! [X2 : 'list_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (~sP8(succ_0(zero_0),cons_0(X1,X2)) | stoogesort_2(cons_0(X1,X3),cons_0(X1,X2)) | ~stoogesort_1(X3,X2)) )),
  inference(resolution,[],[f4945,f887])).
tff(f887,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (sP10(X1,X3,cons_0(X2,nil_0)) | stoogesort_2(cons_0(X2,X0),X3) | ~stoogesort_1(X0,X1)) )),
  inference(resolution,[],[f881,f298])).
tff(f298,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X3 : 'list_0()'] : (~sP9(X4,X2,X0) | stoogesort_2(X0,X3) | sP10(X4,X3,X2)) )),
  inference(cnf_transformation,[],[f298_D])).
tff(f298_D,plain,(
  ( ! [X2,X3,X4] : (( ! [X0] : (~sP9(X4,X2,X0) | stoogesort_2(X0,X3)) ) <=> ~sP10(X4,X3,X2)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP10])])).
tff(f881,plain,(
  ( ! [X4 : 'list_0()',X5 : 'Nat_0()',X3 : 'list_0()'] : (sP9(X4,cons_0(X5,nil_0),cons_0(X5,X3)) | ~stoogesort_1(X3,X4)) )),
  inference(resolution,[],[f873,f296])).
tff(f296,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()'] : (~x_14(X0,X2,X6) | ~stoogesort_1(X6,X4) | sP9(X4,X2,X0)) )),
  inference(cnf_transformation,[],[f296_D])).
tff(f296_D,plain,(
  ( ! [X0,X2,X4] : (( ! [X6] : (~x_14(X0,X2,X6) | ~stoogesort_1(X6,X4)) ) <=> ~sP9(X4,X2,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP9])])).
tff(f873,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (x_14(cons_0(X0,X1),cons_0(X0,nil_0),X1)) )),
  inference(unit_resulting_resolution,[],[f221,f260])).
tff(f260,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (x_14(cons_0(X3,X2),cons_0(X3,X0),X1) | ~x_14(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f162])).
tff(f162,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (x_14(cons_0(X3,X2),cons_0(X3,X0),X1) | ~x_14(X2,X0,X1))),
  inference(ennf_transformation,[],[f117])).
tff(f117,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (x_14(X2,X0,X1) => x_14(cons_0(X3,X2),cons_0(X3,X0),X1))),
  inference(rectify,[],[f69])).
tff(f69,axiom,(
  ! [X2 : 'list_0()',X3 : 'list_0()',X0 : 'list_0()',X1 : 'Nat_0()'] : (x_14(X0,X2,X3) => x_14(cons_0(X1,X0),cons_0(X1,X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f221,plain,(
  ( ! [X0 : 'list_0()'] : (x_14(X0,nil_0,X0)) )),
  inference(cnf_transformation,[],[f70])).
tff(f70,axiom,(
  ! [X0 : 'list_0()'] : x_14(X0,nil_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f4945,plain,(
  ( ! [X2 : 'Nat_0()',X3 : 'list_0()'] : (~sP10(X3,cons_0(X2,X3),cons_0(X2,nil_0)) | ~sP8(succ_0(zero_0),cons_0(X2,X3))) )),
  inference(resolution,[],[f4898,f299])).
tff(f299,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X5 : 'Nat_0()',X3 : 'list_0()'] : (~splitAt_0(pair_1(X2,X4),X5,X3) | ~sP8(X5,X3) | ~sP10(X4,X3,X2)) )),
  inference(general_splitting,[],[f297,f298_D])).
tff(f297,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X5 : 'Nat_0()',X3 : 'list_0()'] : (stoogesort_2(X0,X3) | ~splitAt_0(pair_1(X2,X4),X5,X3) | ~sP8(X5,X3) | ~sP9(X4,X2,X0)) )),
  inference(general_splitting,[],[f295,f296_D])).
tff(f295,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X5 : 'Nat_0()',X3 : 'list_0()'] : (stoogesort_2(X0,X3) | ~stoogesort_1(X6,X4) | ~x_14(X0,X2,X6) | ~splitAt_0(pair_1(X2,X4),X5,X3) | ~sP8(X5,X3)) )),
  inference(general_splitting,[],[f274,f294_D])).
tff(f294,plain,(
  ( ! [X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (~idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) | ~length_0(X1,X3) | sP8(X5,X3)) )),
  inference(cnf_transformation,[],[f294_D])).
tff(f294_D,plain,(
  ( ! [X3,X5] : (( ! [X1] : (~idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) | ~length_0(X1,X3)) ) <=> ~sP8(X5,X3)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP8])])).
tff(f274,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (stoogesort_2(X0,X3) | ~stoogesort_1(X6,X4) | ~x_14(X0,X2,X6) | ~length_0(X1,X3) | ~idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X2,X4),X5,X3)) )),
  inference(cnf_transformation,[],[f184])).
tff(f184,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_2(X0,X3) | ~stoogesort_1(X6,X4) | ~x_14(X0,X2,X6) | ~length_0(X1,X3) | ~idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X2,X4),X5,X3))),
  inference(flattening,[],[f183])).
tff(f183,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_2(X0,X3) | (~stoogesort_1(X6,X4) | ~x_14(X0,X2,X6) | ~length_0(X1,X3) | ~idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X2,X4),X5,X3)))),
  inference(ennf_transformation,[],[f134])).
tff(f134,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : ((stoogesort_1(X6,X4) & x_14(X0,X2,X6) & length_0(X1,X3) & idiv_0(X5,X1,succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X2,X4),X5,X3)) => stoogesort_2(X0,X3))),
  inference(rectify,[],[f76])).
tff(f76,axiom,(
  ! [X0 : 'list_0()',X2 : 'Nat_0()',X4 : 'list_0()',X6 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : ((stoogesort_1(X1,X5) & x_14(X0,X4,X1) & length_0(X2,X6) & idiv_0(X3,X2,succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X4,X5),X3,X6)) => stoogesort_2(X0,X6))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f4898,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (splitAt_0(pair_1(cons_0(X0,nil_0),X1),succ_0(zero_0),cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f517,f1100,f267])).
tff(f267,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (splitAt_0(pair_1(X1,X0),X3,X2) | ~take_0(X1,X3,X2) | ~drop_0(X0,X3,X2)) )),
  inference(cnf_transformation,[],[f172])).
tff(f172,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (splitAt_0(pair_1(X1,X0),X3,X2) | ~take_0(X1,X3,X2) | ~drop_0(X0,X3,X2))),
  inference(flattening,[],[f171])).
tff(f171,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (splitAt_0(pair_1(X1,X0),X3,X2) | (~take_0(X1,X3,X2) | ~drop_0(X0,X3,X2)))),
  inference(ennf_transformation,[],[f125])).
tff(f125,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : ((take_0(X1,X3,X2) & drop_0(X0,X3,X2)) => splitAt_0(pair_1(X1,X0),X3,X2))),
  inference(rectify,[],[f68])).
tff(f68,axiom,(
  ! [X1 : 'list_0()',X0 : 'list_0()',X3 : 'list_0()',X2 : 'Nat_0()'] : ((take_0(X0,X2,X3) & drop_0(X1,X2,X3)) => splitAt_0(pair_1(X0,X1),X2,X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f1100,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (take_0(cons_0(X0,nil_0),succ_0(zero_0),cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f217,f262])).
tff(f262,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (take_0(cons_0(X0,X1),succ_0(X2),cons_0(X0,X3)) | ~take_0(X1,X2,X3)) )),
  inference(cnf_transformation,[],[f164])).
tff(f164,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (take_0(cons_0(X0,X1),succ_0(X2),cons_0(X0,X3)) | ~take_0(X1,X2,X3))),
  inference(ennf_transformation,[],[f119])).
tff(f119,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (take_0(X1,X2,X3) => take_0(cons_0(X0,X1),succ_0(X2),cons_0(X0,X3)))),
  inference(rectify,[],[f38])).
tff(f38,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'list_0()',X3 : 'Nat_0()',X2 : 'list_0()'] : (take_0(X0,X3,X2) => take_0(cons_0(X1,X0),succ_0(X3),cons_0(X1,X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f217,plain,(
  ( ! [X0 : 'list_0()'] : (take_0(nil_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f40])).
tff(f40,axiom,(
  ! [X0 : 'list_0()'] : take_0(nil_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f517,plain,(
  ( ! [X0 : 'list_0()',X1 : 'Nat_0()'] : (drop_0(X0,succ_0(zero_0),cons_0(X1,X0))) )),
  inference(unit_resulting_resolution,[],[f223,f261])).
tff(f261,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (drop_0(X0,succ_0(X1),cons_0(X2,X3)) | ~drop_0(X0,X1,X3)) )),
  inference(cnf_transformation,[],[f163])).
tff(f163,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (drop_0(X0,succ_0(X1),cons_0(X2,X3)) | ~drop_0(X0,X1,X3))),
  inference(ennf_transformation,[],[f118])).
tff(f118,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (drop_0(X0,X1,X3) => drop_0(X0,succ_0(X1),cons_0(X2,X3)))),
  inference(rectify,[],[f65])).
tff(f65,axiom,(
  ! [X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()'] : (drop_0(X0,X3,X2) => drop_0(X0,succ_0(X3),cons_0(X1,X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f223,plain,(
  ( ! [X0 : 'list_0()'] : (drop_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f67])).
tff(f67,axiom,(
  ! [X0 : 'list_0()'] : drop_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f7742,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP8(succ_0(zero_0),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f679,f7728,f294])).
tff(f7728,plain,(
  idiv_0(succ_0(zero_0),succ_0(succ_0(succ_0(zero_0))),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f344,f511,f407,f281])).
tff(f281,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~idiv_0(X1,X2,X0) | ~minus_0(X2,X4,X0) | idiv_0(succ_0(X1),X4,X0) | ~sP1(X4,X0)) )),
  inference(general_splitting,[],[f269,f280_D])).
tff(f280,plain,(
  ( ! [X4 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Bool_0()'] : (~lt_0(X3,X4,X0) | ~diseqBool_0(X3,true_0) | sP1(X4,X0)) )),
  inference(cnf_transformation,[],[f280_D])).
tff(f280_D,plain,(
  ( ! [X0,X4] : (( ! [X3] : (~lt_0(X3,X4,X0) | ~diseqBool_0(X3,true_0)) ) <=> ~sP1(X4,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP1])])).
tff(f269,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Bool_0()',X1 : 'Nat_0()'] : (idiv_0(succ_0(X1),X4,X0) | ~diseqBool_0(X3,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X1,X2,X0) | ~lt_0(X3,X4,X0)) )),
  inference(cnf_transformation,[],[f176])).
tff(f176,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : (idiv_0(succ_0(X1),X4,X0) | ~diseqBool_0(X3,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X1,X2,X0) | ~lt_0(X3,X4,X0))),
  inference(flattening,[],[f175])).
tff(f175,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : (idiv_0(succ_0(X1),X4,X0) | (~diseqBool_0(X3,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X1,X2,X0) | ~lt_0(X3,X4,X0)))),
  inference(ennf_transformation,[],[f127])).
tff(f127,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Bool_0()',X4 : 'Nat_0()'] : ((diseqBool_0(X3,true_0) & minus_0(X2,X4,X0) & idiv_0(X1,X2,X0) & lt_0(X3,X4,X0)) => idiv_0(succ_0(X1),X4,X0))),
  inference(rectify,[],[f64])).
tff(f64,axiom,(
  ! [X4 : 'Nat_0()',X1 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Bool_0()',X3 : 'Nat_0()'] : ((diseqBool_0(X2,true_0) & minus_0(X0,X3,X4) & idiv_0(X1,X0,X4) & lt_0(X2,X3,X4)) => idiv_0(succ_0(X1),X3,X4))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f407,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(zero_0))),succ_0(succ_0(succ_0(X0))))) )),
  inference(unit_resulting_resolution,[],[f405,f251])).
tff(f251,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (minus_0(X1,succ_0(X2),succ_0(X0)) | ~minus_0(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f152])).
tff(f152,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X1,succ_0(X2),succ_0(X0)) | ~minus_0(X1,X2,X0))),
  inference(ennf_transformation,[],[f108])).
tff(f108,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X1,X2,X0) => minus_0(X1,succ_0(X2),succ_0(X0)))),
  inference(rectify,[],[f45])).
tff(f45,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X0,X2,X1) => minus_0(X0,succ_0(X2),succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f405,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(zero_0)),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f403,f251])).
tff(f403,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(zero_0),succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f218,f251])).
tff(f218,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f47])).
tff(f47,axiom,(
  ! [X0 : 'Nat_0()'] : minus_0(zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f511,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP1(succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f196,f416,f280])).
tff(f416,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f412,f252])).
tff(f252,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Bool_0()',X1 : 'Nat_0()'] : (lt_0(X0,succ_0(X1),succ_0(X2)) | ~lt_0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f153])).
tff(f153,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (lt_0(X0,succ_0(X1),succ_0(X2)) | ~lt_0(X0,X1,X2))),
  inference(ennf_transformation,[],[f48])).
tff(f48,axiom,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (lt_0(X0,X1,X2) => lt_0(X0,succ_0(X1),succ_0(X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f412,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(succ_0(X0)),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f409,f252])).
tff(f409,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(X0),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f219,f252])).
tff(f219,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,X0,zero_0)) )),
  inference(cnf_transformation,[],[f50])).
tff(f50,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(false_0,X0,zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f196,plain,(
  diseqBool_0(false_0,true_0)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  diseqBool_0(false_0,true_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f344,plain,(
  ( ! [X0 : 'Nat_0()'] : (idiv_0(zero_0,zero_0,succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f227,f240])).
tff(f240,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~lt_0(true_0,X1,X0) | idiv_0(zero_0,X1,X0)) )),
  inference(cnf_transformation,[],[f141])).
tff(f141,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (idiv_0(zero_0,X1,X0) | ~lt_0(true_0,X1,X0))),
  inference(ennf_transformation,[],[f94])).
tff(f94,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (lt_0(true_0,X1,X0) => idiv_0(zero_0,X1,X0))),
  inference(rectify,[],[f63])).
tff(f63,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : (lt_0(true_0,X0,X1) => idiv_0(zero_0,X0,X1))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f227,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(true_0,zero_0,succ_0(X0))) )),
  inference(cnf_transformation,[],[f49])).
tff(f49,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(true_0,zero_0,succ_0(X0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f679,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(succ_0(succ_0(succ_0(zero_0))),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f675,f674])).
tff(f674,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (length_0(succ_0(X0),cons_0(X2,X1)) | ~length_0(X0,X1)) )),
  inference(resolution,[],[f668,f279])).
tff(f279,plain,(
  ( ! [X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~sP0(X1,X0) | length_0(X1,cons_0(X3,X0))) )),
  inference(general_splitting,[],[f268,f278_D])).
tff(f278,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X1 : 'Nat_0()'] : (~plus_0(X1,succ_0(zero_0),X2) | ~length_0(X2,X0) | sP0(X1,X0)) )),
  inference(cnf_transformation,[],[f278_D])).
tff(f278_D,plain,(
  ( ! [X0,X1] : (( ! [X2] : (~plus_0(X1,succ_0(zero_0),X2) | ~length_0(X2,X0)) ) <=> ~sP0(X1,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP0])])).
tff(f268,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(X1,cons_0(X3,X0)) | ~length_0(X2,X0) | ~plus_0(X1,succ_0(zero_0),X2)) )),
  inference(cnf_transformation,[],[f174])).
tff(f174,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (length_0(X1,cons_0(X3,X0)) | ~length_0(X2,X0) | ~plus_0(X1,succ_0(zero_0),X2))),
  inference(flattening,[],[f173])).
tff(f173,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (length_0(X1,cons_0(X3,X0)) | (~length_0(X2,X0) | ~plus_0(X1,succ_0(zero_0),X2)))),
  inference(ennf_transformation,[],[f126])).
tff(f126,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : ((length_0(X2,X0) & plus_0(X1,succ_0(zero_0),X2)) => length_0(X1,cons_0(X3,X0)))),
  inference(rectify,[],[f56])).
tff(f56,axiom,(
  ! [X3 : 'list_0()',X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : ((length_0(X1,X3) & plus_0(X0,succ_0(zero_0),X1)) => length_0(X0,cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f668,plain,(
  ( ! [X4 : 'list_0()',X3 : 'Nat_0()'] : (sP0(succ_0(X3),X4) | ~length_0(X3,X4)) )),
  inference(resolution,[],[f278,f375])).
tff(f375,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f222,f250])).
tff(f250,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (plus_0(succ_0(X2),succ_0(X0),X1) | ~plus_0(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f151])).
tff(f151,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(succ_0(X2),succ_0(X0),X1) | ~plus_0(X2,X0,X1))),
  inference(ennf_transformation,[],[f107])).
tff(f107,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X2,X0,X1) => plus_0(succ_0(X2),succ_0(X0),X1))),
  inference(rectify,[],[f41])).
tff(f41,axiom,(
  ! [X1 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()'] : (plus_0(X0,X1,X2) => plus_0(succ_0(X0),succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f222,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f42])).
tff(f42,axiom,(
  ! [X0 : 'Nat_0()'] : plus_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f675,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(succ_0(succ_0(zero_0)),cons_0(X0,cons_0(X1,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f671,f279])).
tff(f671,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP0(succ_0(succ_0(zero_0)),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f375,f669,f278])).
tff(f669,plain,(
  ( ! [X0 : 'Nat_0()'] : (length_0(succ_0(zero_0),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f666,f279])).
tff(f666,plain,(
  sP0(succ_0(zero_0),nil_0)),
  inference(unit_resulting_resolution,[],[f197,f375,f278])).
tff(f197,plain,(
  length_0(zero_0,nil_0)),
  inference(cnf_transformation,[],[f57])).
tff(f57,axiom,(
  length_0(zero_0,nil_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f852,plain,(
  ( ! [X0 : 'Nat_0()'] : (stoogesort_1(cons_0(zero_0,cons_0(X0,nil_0)),cons_0(zero_0,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f848,f249])).
tff(f249,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_1(X2,cons_0(X0,cons_0(X1,nil_0))) | ~sort_0(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f150])).
tff(f150,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()'] : (stoogesort_1(X2,cons_0(X0,cons_0(X1,nil_0))) | ~sort_0(X2,X0,X1))),
  inference(ennf_transformation,[],[f106])).
tff(f106,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()'] : (sort_0(X2,X0,X1) => stoogesort_1(X2,cons_0(X0,cons_0(X1,nil_0))))),
  inference(rectify,[],[f73])).
tff(f73,axiom,(
  ! [X2 : 'Nat_0()',X1 : 'Nat_0()',X0 : 'list_0()'] : (sort_0(X0,X2,X1) => stoogesort_1(X0,cons_0(X2,cons_0(X1,nil_0))))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f848,plain,(
  ( ! [X0 : 'Nat_0()'] : (sort_0(cons_0(zero_0,cons_0(X0,nil_0)),zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f220,f241])).
tff(f241,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sort_0(cons_0(X1,cons_0(X0,nil_0)),X1,X0) | ~leq_0(true_0,X1,X0)) )),
  inference(cnf_transformation,[],[f142])).
tff(f142,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sort_0(cons_0(X1,cons_0(X0,nil_0)),X1,X0) | ~leq_0(true_0,X1,X0))),
  inference(ennf_transformation,[],[f95])).
tff(f95,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (leq_0(true_0,X1,X0) => sort_0(cons_0(X1,cons_0(X0,nil_0)),X1,X0))),
  inference(rectify,[],[f54])).
tff(f54,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : (leq_0(true_0,X0,X1) => sort_0(cons_0(X0,cons_0(X1,nil_0)),X0,X1))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f220,plain,(
  ( ! [X0 : 'Nat_0()'] : (leq_0(true_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f53])).
tff(f53,axiom,(
  ! [X0 : 'Nat_0()'] : leq_0(true_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f13202,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_0(cons_0(X0,cons_0(X1,cons_0(X2,nil_0))),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f13167,f8239])).
tff(f8239,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (~sP18(cons_0(X0,X1),succ_0(zero_0)) | stoogesort_0(cons_0(X0,X1),cons_0(X0,X1))) )),
  inference(resolution,[],[f4944,f8237])).
tff(f8237,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (sP20(X1,cons_0(X0,nil_0),X2) | stoogesort_0(cons_0(X0,X1),X2)) )),
  inference(resolution,[],[f883,f230])).
tff(f230,plain,(
  ( ! [X0 : 'Nat_0()'] : (stoogesort_1(cons_0(X0,nil_0),cons_0(X0,nil_0))) )),
  inference(cnf_transformation,[],[f74])).
tff(f74,axiom,(
  ! [X0 : 'Nat_0()'] : stoogesort_1(cons_0(X0,nil_0),cons_0(X0,nil_0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f883,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~stoogesort_1(cons_0(X0,nil_0),X3) | stoogesort_0(cons_0(X0,X1),X2) | sP20(X1,X3,X2)) )),
  inference(resolution,[],[f880,f318])).
tff(f318,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~sP19(X3,X2,X0) | ~stoogesort_1(X2,X1) | sP20(X3,X1,X0)) )),
  inference(cnf_transformation,[],[f318_D])).
tff(f318_D,plain,(
  ( ! [X0,X1,X3] : (( ! [X2] : (~sP19(X3,X2,X0) | ~stoogesort_1(X2,X1)) ) <=> ~sP20(X3,X1,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP20])])).
tff(f880,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (sP19(X1,cons_0(X0,nil_0),X2) | stoogesort_0(cons_0(X0,X1),X2)) )),
  inference(resolution,[],[f873,f316])).
tff(f316,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X3 : 'list_0()'] : (~x_14(X4,X2,X3) | stoogesort_0(X4,X0) | sP19(X3,X2,X0)) )),
  inference(cnf_transformation,[],[f316_D])).
tff(f316_D,plain,(
  ( ! [X0,X2,X3] : (( ! [X4] : (~x_14(X4,X2,X3) | stoogesort_0(X4,X0)) ) <=> ~sP19(X3,X2,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP19])])).
tff(f4944,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (~sP20(X1,cons_0(X0,nil_0),cons_0(X0,X1)) | ~sP18(cons_0(X0,X1),succ_0(zero_0))) )),
  inference(resolution,[],[f4898,f319])).
tff(f319,plain,(
  ( ! [X0 : 'list_0()',X7 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~splitAt_0(pair_1(X1,X3),X7,X0) | ~sP18(X0,X7) | ~sP20(X3,X1,X0)) )),
  inference(general_splitting,[],[f317,f318_D])).
tff(f317,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X7 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~stoogesort_1(X2,X1) | ~splitAt_0(pair_1(X1,X3),X7,X0) | ~sP18(X0,X7) | ~sP19(X3,X2,X0)) )),
  inference(general_splitting,[],[f315,f316_D])).
tff(f315,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X7 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X4,X0) | ~stoogesort_1(X2,X1) | ~x_14(X4,X2,X3) | ~splitAt_0(pair_1(X1,X3),X7,X0) | ~sP18(X0,X7)) )),
  inference(general_splitting,[],[f313,f314_D])).
tff(f314,plain,(
  ( ! [X0 : 'list_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~sP17(X5,X7) | ~length_0(X5,X0) | sP18(X0,X7)) )),
  inference(cnf_transformation,[],[f314_D])).
tff(f314_D,plain,(
  ( ! [X7,X0] : (( ! [X5] : (~sP17(X5,X7) | ~length_0(X5,X0)) ) <=> ~sP18(X0,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP18])])).
tff(f313,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X4,X0) | ~stoogesort_1(X2,X1) | ~x_14(X4,X2,X3) | ~length_0(X5,X0) | ~splitAt_0(pair_1(X1,X3),X7,X0) | ~sP17(X5,X7)) )),
  inference(general_splitting,[],[f277,f312_D])).
tff(f312,plain,(
  ( ! [X6 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~times_0(X6,succ_0(succ_0(zero_0)),X5) | sP17(X5,X7)) )),
  inference(cnf_transformation,[],[f312_D])).
tff(f312_D,plain,(
  ( ! [X7,X5] : (( ! [X6] : (~idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~times_0(X6,succ_0(succ_0(zero_0)),X5)) ) <=> ~sP17(X5,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP17])])).
tff(f277,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X4,X0) | ~stoogesort_1(X2,X1) | ~x_14(X4,X2,X3) | ~length_0(X5,X0) | ~times_0(X6,succ_0(succ_0(zero_0)),X5) | ~idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X1,X3),X7,X0)) )),
  inference(cnf_transformation,[],[f188])).
tff(f188,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (stoogesort_0(X4,X0) | ~stoogesort_1(X2,X1) | ~x_14(X4,X2,X3) | ~length_0(X5,X0) | ~times_0(X6,succ_0(succ_0(zero_0)),X5) | ~idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X1,X3),X7,X0))),
  inference(flattening,[],[f187])).
tff(f187,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (stoogesort_0(X4,X0) | (~stoogesort_1(X2,X1) | ~x_14(X4,X2,X3) | ~length_0(X5,X0) | ~times_0(X6,succ_0(succ_0(zero_0)),X5) | ~idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X1,X3),X7,X0)))),
  inference(ennf_transformation,[],[f139])).
tff(f139,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : ((stoogesort_1(X2,X1) & x_14(X4,X2,X3) & length_0(X5,X0) & times_0(X6,succ_0(succ_0(zero_0)),X5) & idiv_0(X7,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X1,X3),X7,X0)) => stoogesort_0(X4,X0))),
  inference(rectify,[],[f71])).
tff(f71,axiom,(
  ! [X7 : 'list_0()',X5 : 'list_0()',X0 : 'list_0()',X6 : 'list_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : ((stoogesort_1(X0,X5) & x_14(X1,X0,X6) & length_0(X2,X7) & times_0(X3,succ_0(succ_0(zero_0)),X2) & idiv_0(X4,succ_0(X3),succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X5,X6),X4,X7)) => stoogesort_0(X1,X7))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f13167,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP18(cons_0(X0,cons_0(X1,cons_0(X2,nil_0))),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f679,f13087,f314])).
tff(f13087,plain,(
  sP17(succ_0(succ_0(succ_0(zero_0))),succ_0(zero_0))),
  inference(unit_resulting_resolution,[],[f7729,f1516,f312])).
tff(f1516,plain,(
  times_0(succ_0(succ_0(succ_0(succ_0(succ_0(succ_0(zero_0)))))),succ_0(succ_0(zero_0)),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f389,f1107,f264])).
tff(f264,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~times_0(X1,X2,X3) | times_0(X0,succ_0(X2),X3) | ~plus_0(X0,X3,X1)) )),
  inference(cnf_transformation,[],[f167])).
tff(f167,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X0,succ_0(X2),X3) | ~times_0(X1,X2,X3) | ~plus_0(X0,X3,X1))),
  inference(flattening,[],[f166])).
tff(f166,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X0,succ_0(X2),X3) | (~times_0(X1,X2,X3) | ~plus_0(X0,X3,X1)))),
  inference(ennf_transformation,[],[f43])).
tff(f43,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : ((times_0(X1,X2,X3) & plus_0(X0,X3,X1)) => times_0(X0,succ_0(X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f1107,plain,(
  times_0(succ_0(succ_0(succ_0(zero_0))),succ_0(zero_0),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f389,f216,f264])).
tff(f216,plain,(
  ( ! [X0 : 'Nat_0()'] : (times_0(zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f44])).
tff(f44,axiom,(
  ! [X0 : 'Nat_0()'] : times_0(zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f389,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))),X0)) )),
  inference(unit_resulting_resolution,[],[f386,f250])).
tff(f386,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(succ_0(X0)),succ_0(succ_0(zero_0)),X0)) )),
  inference(unit_resulting_resolution,[],[f375,f250])).
tff(f7729,plain,(
  ( ! [X0 : 'Nat_0()'] : (idiv_0(succ_0(zero_0),succ_0(succ_0(succ_0(succ_0(X0)))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f344,f511,f408,f281])).
tff(f408,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(succ_0(X0)))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f406,f251])).
tff(f406,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f404,f251])).
tff(f404,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(X0)),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f226,f251])).
tff(f226,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f46])).
tff(f46,axiom,(
  ! [X0 : 'Nat_0()'] : minus_0(zero_0,succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f427,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (diseqlist_0(cons_0(succ_0(X0),X1),cons_0(zero_0,X2))) )),
  inference(unit_resulting_resolution,[],[f214,f256])).
tff(f256,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'list_0()',X1 : 'list_0()'] : (diseqlist_0(cons_0(X2,X1),cons_0(X0,X3)) | ~diseqNat_0(X2,X0)) )),
  inference(cnf_transformation,[],[f158])).
tff(f158,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (diseqlist_0(cons_0(X2,X1),cons_0(X0,X3)) | ~diseqNat_0(X2,X0))),
  inference(ennf_transformation,[],[f113])).
tff(f113,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (diseqNat_0(X2,X0) => diseqlist_0(cons_0(X2,X1),cons_0(X0,X3)))),
  inference(rectify,[],[f31])).
tff(f31,axiom,(
  ! [X2 : 'Nat_0()',X1 : 'list_0()',X0 : 'Nat_0()',X3 : 'list_0()'] : (diseqNat_0(X0,X2) => diseqlist_0(cons_0(X0,X1),cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f214,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f23])).
tff(f23,axiom,(
  ! [X0 : 'Nat_0()'] : diseqNat_0(succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f4396,plain,(
  ( ! [X0 : 'Nat_0()'] : (isort_0(cons_0(zero_0,cons_0(X0,nil_0)),cons_0(zero_0,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f888,f4390,f266])).
tff(f4390,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (insert_0(cons_0(zero_0,cons_0(X0,X1)),zero_0,cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f220,f248])).
tff(f248,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X1 : 'Nat_0()'] : (insert_0(cons_0(X2,cons_0(X1,X0)),X2,cons_0(X1,X0)) | ~leq_0(true_0,X2,X1)) )),
  inference(cnf_transformation,[],[f149])).
tff(f149,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (insert_0(cons_0(X2,cons_0(X1,X0)),X2,cons_0(X1,X0)) | ~leq_0(true_0,X2,X1))),
  inference(ennf_transformation,[],[f105])).
tff(f105,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (leq_0(true_0,X2,X1) => insert_0(cons_0(X2,cons_0(X1,X0)),X2,cons_0(X1,X0)))),
  inference(rectify,[],[f58])).
tff(f58,axiom,(
  ! [X1 : 'list_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : (leq_0(true_0,X2,X0) => insert_0(cons_0(X2,cons_0(X0,X1)),X2,cons_0(X0,X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f888,plain,(
  ( ! [X0 : 'Nat_0()'] : (isort_0(cons_0(X0,nil_0),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f229,f195,f266])).
tff(f195,plain,(
  isort_0(nil_0,nil_0)),
  inference(cnf_transformation,[],[f62])).
tff(f62,axiom,(
  isort_0(nil_0,nil_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f229,plain,(
  ( ! [X0 : 'Nat_0()'] : (insert_0(cons_0(X0,nil_0),X0,nil_0)) )),
  inference(cnf_transformation,[],[f60])).
tff(f60,axiom,(
  ! [X0 : 'Nat_0()'] : insert_0(cons_0(X0,nil_0),X0,nil_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f7104,plain,(
  ( ! [X0 : 'Nat_0()'] : (insert_0(cons_0(zero_0,cons_0(zero_0,cons_0(succ_0(X0),nil_0))),succ_0(X0),cons_0(zero_0,cons_0(zero_0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f463,f7092,f283])).
tff(f283,plain,(
  ( ! [X4 : 'Nat_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (~insert_0(X1,X4,X0) | insert_0(cons_0(X3,X1),X4,cons_0(X3,X0)) | ~sP2(X4,X3)) )),
  inference(general_splitting,[],[f270,f282_D])).
tff(f282,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Bool_0()',X3 : 'Nat_0()'] : (~leq_0(X2,X4,X3) | ~diseqBool_0(X2,true_0) | sP2(X4,X3)) )),
  inference(cnf_transformation,[],[f282_D])).
tff(f282_D,plain,(
  ( ! [X3,X4] : (( ! [X2] : (~leq_0(X2,X4,X3) | ~diseqBool_0(X2,true_0)) ) <=> ~sP2(X4,X3)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP2])])).
tff(f270,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Bool_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (insert_0(cons_0(X3,X1),X4,cons_0(X3,X0)) | ~diseqBool_0(X2,true_0) | ~insert_0(X1,X4,X0) | ~leq_0(X2,X4,X3)) )),
  inference(cnf_transformation,[],[f178])).
tff(f178,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'Bool_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : (insert_0(cons_0(X3,X1),X4,cons_0(X3,X0)) | ~diseqBool_0(X2,true_0) | ~insert_0(X1,X4,X0) | ~leq_0(X2,X4,X3))),
  inference(flattening,[],[f177])).
tff(f177,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'Bool_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : (insert_0(cons_0(X3,X1),X4,cons_0(X3,X0)) | (~diseqBool_0(X2,true_0) | ~insert_0(X1,X4,X0) | ~leq_0(X2,X4,X3)))),
  inference(ennf_transformation,[],[f128])).
tff(f128,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'Bool_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : ((diseqBool_0(X2,true_0) & insert_0(X1,X4,X0) & leq_0(X2,X4,X3)) => insert_0(cons_0(X3,X1),X4,cons_0(X3,X0)))),
  inference(rectify,[],[f59])).
tff(f59,axiom,(
  ! [X3 : 'list_0()',X0 : 'list_0()',X1 : 'Bool_0()',X2 : 'Nat_0()',X4 : 'Nat_0()'] : ((diseqBool_0(X1,true_0) & insert_0(X0,X4,X3) & leq_0(X1,X4,X2)) => insert_0(cons_0(X2,X0),X4,cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
tff(f7092,plain,(
  ( ! [X0 : 'Nat_0()'] : (insert_0(cons_0(zero_0,cons_0(succ_0(X0),nil_0)),succ_0(X0),cons_0(zero_0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f463,f229,f283])).
tff(f463,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP2(succ_0(X0),zero_0)) )),
  inference(unit_resulting_resolution,[],[f228,f196,f282])).
tff(f228,plain,(
  ( ! [X0 : 'Nat_0()'] : (leq_0(false_0,succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f52])).
tff(f52,axiom,(
  ! [X0 : 'Nat_0()'] : leq_0(false_0,succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2IsSort.smt2',unknown)).
