11
22472
650
UNSAT
tff(type_def_5, type, 'Bool_0()': $tType).
tff(type_def_6, type, 'Nat_0()': $tType).
tff(type_def_7, type, 'list_0()': $tType).
tff(type_def_8, type, 'pair_0()': $tType).
tff(func_def_0, type, false_0: 'Bool_0()').
tff(func_def_1, type, true_0: 'Bool_0()').
tff(func_def_2, type, zero_0: 'Nat_0()').
tff(func_def_3, type, succ_0: ('Nat_0()') > 'Nat_0()').
tff(func_def_4, type, nil_0: 'list_0()').
tff(func_def_5, type, cons_0: ('Nat_0()' * 'list_0()') > 'list_0()').
tff(func_def_6, type, pair_1: ('list_0()' * 'list_0()') > 'pair_0()').
tff(pred_def_1, type, diseqBool_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_2, type, isfalse_1: ('Bool_0()') > $o).
tff(pred_def_3, type, istrue_1: ('Bool_0()') > $o).
tff(pred_def_4, type, and_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_5, type, or_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_6, type, hence_0: ('Bool_0()' * 'Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_7, type, not_0: ('Bool_0()' * 'Bool_0()') > $o).
tff(pred_def_8, type, diseqNat_0: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_9, type, p_1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_10, type, iszero_0: ('Nat_0()') > $o).
tff(pred_def_11, type, issucc_0: ('Nat_0()') > $o).
tff(pred_def_12, type, diseqlist_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_13, type, head_1: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_14, type, tail_1: ('list_0()' * 'list_0()') > $o).
tff(pred_def_15, type, isnil_0: ('list_0()') > $o).
tff(pred_def_16, type, iscons_0: ('list_0()') > $o).
tff(pred_def_17, type, diseqpair_0: ('pair_0()' * 'pair_0()') > $o).
tff(pred_def_18, type, projpair_2: ('list_0()' * 'pair_0()') > $o).
tff(pred_def_19, type, projpair_3: ('list_0()' * 'pair_0()') > $o).
tff(pred_def_20, type, ispair_0: ('pair_0()') > $o).
tff(pred_def_21, type, take_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_22, type, plus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_23, type, times_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_24, type, minus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_25, type, lt_0: ('Bool_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_26, type, leq_0: ('Bool_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_27, type, ordered_0: ('Bool_0()' * 'list_0()') > $o).
tff(pred_def_28, type, sort_0: ('list_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_29, type, length_0: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_30, type, idiv_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_31, type, drop_0: ('list_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_32, type, splitAt_0: ('pair_0()' * 'Nat_0()' * 'list_0()') > $o).
tff(pred_def_33, type, x_13: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_34, type, stoogesort_0: ('list_0()' * 'list_0()') > $o).
tff(pred_def_35, type, stoogesort_1: ('list_0()' * 'list_0()') > $o).
tff(pred_def_36, type, stoogesort_2: ('list_0()' * 'list_0()') > $o).
tff(pred_def_37, type, sP0: ('list_0()') > $o).
tff(pred_def_38, type, sP1: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_39, type, sP2: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_40, type, sP3: ('Bool_0()' * 'Nat_0()' * 'Bool_0()' * 'Nat_0()') > $o).
tff(pred_def_41, type, sP4: ('list_0()' * 'list_0()') > $o).
tff(pred_def_42, type, sP5: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_43, type, sP6: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_44, type, sP7: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_45, type, sP8: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_46, type, sP9: ('Nat_0()' * 'list_0()') > $o).
tff(pred_def_47, type, sP10: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_48, type, sP11: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_49, type, sP12: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_50, type, sP13: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_51, type, sP14: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_52, type, sP15: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_53, type, sP16: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_54, type, sP17: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_55, type, sP18: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_56, type, sP19: ('list_0()' * 'Nat_0()') > $o).
tff(pred_def_57, type, sP20: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(pred_def_58, type, sP21: ('list_0()' * 'list_0()' * 'list_0()') > $o).
tff(f34048,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f33969,f7271,f34036,f278])).
tff(f278,plain,(
  ( ! [X6 : 'list_0()',X4 : 'Nat_0()',X2 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~stoogesort_0(X6,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | stoogesort_1(X5,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~sP4(X5,X6)) )),
  inference(general_splitting,[],[f262,f277_D])).
tff(f277,plain,(
  ( ! [X6 : 'list_0()',X0 : 'list_0()',X5 : 'list_0()'] : (~stoogesort_2(X0,X6) | ~stoogesort_0(X5,X0) | sP4(X5,X6)) )),
  inference(cnf_transformation,[],[f277_D])).
tff(f277_D,plain,(
  ( ! [X6,X5] : (( ! [X0] : (~stoogesort_2(X0,X6) | ~stoogesort_0(X5,X0)) ) <=> ~sP4(X5,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP4])])).
tff(f262,plain,(
  ( ! [X6 : 'list_0()',X4 : 'Nat_0()',X2 : 'list_0()',X0 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_1(X5,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~stoogesort_0(X6,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~stoogesort_2(X0,X6) | ~stoogesort_0(X5,X0)) )),
  inference(cnf_transformation,[],[f173])).
tff(f173,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'list_0()',X6 : 'list_0()'] : (stoogesort_1(X5,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~stoogesort_0(X6,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~stoogesort_2(X0,X6) | ~stoogesort_0(X5,X0))),
  inference(flattening,[],[f172])).
tff(f172,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'list_0()',X6 : 'list_0()'] : (stoogesort_1(X5,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | (~stoogesort_0(X6,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) | ~stoogesort_2(X0,X6) | ~stoogesort_0(X5,X0)))),
  inference(ennf_transformation,[],[f125])).
tff(f125,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'list_0()',X6 : 'list_0()'] : ((stoogesort_0(X6,cons_0(X4,cons_0(X3,cons_0(X1,X2)))) & stoogesort_2(X0,X6) & stoogesort_0(X5,X0)) => stoogesort_1(X5,cons_0(X4,cons_0(X3,cons_0(X1,X2)))))),
  inference(rectify,[],[f70])).
tff(f70,axiom,(
  ! [X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X0 : 'list_0()',X1 : 'list_0()'] : ((stoogesort_0(X1,cons_0(X6,cons_0(X5,cons_0(X3,X4)))) & stoogesort_2(X2,X1) & stoogesort_0(X0,X2)) => stoogesort_1(X0,cons_0(X6,cons_0(X5,cons_0(X3,X4)))))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f34036,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP4(cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f33803,f33969,f277])).
tff(f33803,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_2(cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f853,f6658,f292])).
tff(f292,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X0 : 'list_0()',X1 : 'list_0()'] : (~sP11(X4,X0,X6) | ~sP10(X4,X1,X6) | stoogesort_2(X1,X0)) )),
  inference(general_splitting,[],[f290,f291_D])).
tff(f291,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()'] : (~splitAt_0(pair_1(X4,X6),X3,X0) | ~sP9(X3,X0) | sP11(X4,X0,X6)) )),
  inference(cnf_transformation,[],[f291_D])).
tff(f291_D,plain,(
  ( ! [X6,X0,X4] : (( ! [X3] : (~splitAt_0(pair_1(X4,X6),X3,X0) | ~sP9(X3,X0)) ) <=> ~sP11(X4,X0,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP11])])).
tff(f290,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (stoogesort_2(X1,X0) | ~splitAt_0(pair_1(X4,X6),X3,X0) | ~sP9(X3,X0) | ~sP10(X4,X1,X6)) )),
  inference(general_splitting,[],[f288,f289_D])).
tff(f289,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X1 : 'list_0()'] : (~x_13(X1,X4,X2) | ~stoogesort_1(X2,X6) | sP10(X4,X1,X6)) )),
  inference(cnf_transformation,[],[f289_D])).
tff(f289_D,plain,(
  ( ! [X6,X1,X4] : (( ! [X2] : (~x_13(X1,X4,X2) | ~stoogesort_1(X2,X6)) ) <=> ~sP10(X4,X1,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP10])])).
tff(f288,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (stoogesort_2(X1,X0) | ~stoogesort_1(X2,X6) | ~x_13(X1,X4,X2) | ~splitAt_0(pair_1(X4,X6),X3,X0) | ~sP9(X3,X0)) )),
  inference(general_splitting,[],[f265,f287_D])).
tff(f287,plain,(
  ( ! [X0 : 'list_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (~idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) | ~length_0(X5,X0) | sP9(X3,X0)) )),
  inference(cnf_transformation,[],[f287_D])).
tff(f287_D,plain,(
  ( ! [X0,X3] : (( ! [X5] : (~idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) | ~length_0(X5,X0)) ) <=> ~sP9(X3,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP9])])).
tff(f265,plain,(
  ( ! [X6 : 'list_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'list_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (stoogesort_2(X1,X0) | ~stoogesort_1(X2,X6) | ~x_13(X1,X4,X2) | ~length_0(X5,X0) | ~idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X4,X6),X3,X0)) )),
  inference(cnf_transformation,[],[f177])).
tff(f177,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_2(X1,X0) | ~stoogesort_1(X2,X6) | ~x_13(X1,X4,X2) | ~length_0(X5,X0) | ~idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X4,X6),X3,X0))),
  inference(flattening,[],[f176])).
tff(f176,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : (stoogesort_2(X1,X0) | (~stoogesort_1(X2,X6) | ~x_13(X1,X4,X2) | ~length_0(X5,X0) | ~idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X4,X6),X3,X0)))),
  inference(ennf_transformation,[],[f130])).
tff(f130,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()',X5 : 'Nat_0()',X6 : 'list_0()'] : ((stoogesort_1(X2,X6) & x_13(X1,X4,X2) & length_0(X5,X0) & idiv_0(X3,X5,succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X4,X6),X3,X0)) => stoogesort_2(X1,X0))),
  inference(rectify,[],[f74])).
tff(f74,axiom,(
  ! [X6 : 'list_0()',X0 : 'list_0()',X1 : 'list_0()',X3 : 'Nat_0()',X4 : 'list_0()',X2 : 'Nat_0()',X5 : 'list_0()'] : ((stoogesort_1(X1,X5) & x_13(X0,X4,X1) & length_0(X2,X6) & idiv_0(X3,X2,succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X4,X5),X3,X6)) => stoogesort_2(X0,X6))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f6658,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP11(cons_0(X0,nil_0),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))),cons_0(X1,cons_0(X2,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f3897,f6652,f291])).
tff(f6652,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP9(succ_0(zero_0),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f648,f6638,f287])).
tff(f6638,plain,(
  idiv_0(succ_0(zero_0),succ_0(succ_0(succ_0(zero_0))),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f341,f485,f401,f274])).
tff(f274,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()'] : (~idiv_0(X3,X2,X0) | ~minus_0(X2,X4,X0) | idiv_0(succ_0(X3),X4,X0) | ~sP2(X4,X0)) )),
  inference(general_splitting,[],[f260,f273_D])).
tff(f273,plain,(
  ( ! [X4 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Bool_0()'] : (~lt_0(X1,X4,X0) | ~diseqBool_0(X1,true_0) | sP2(X4,X0)) )),
  inference(cnf_transformation,[],[f273_D])).
tff(f273_D,plain,(
  ( ! [X0,X4] : (( ! [X1] : (~lt_0(X1,X4,X0) | ~diseqBool_0(X1,true_0)) ) <=> ~sP2(X4,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP2])])).
tff(f260,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Bool_0()'] : (idiv_0(succ_0(X3),X4,X0) | ~diseqBool_0(X1,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X3,X2,X0) | ~lt_0(X1,X4,X0)) )),
  inference(cnf_transformation,[],[f169])).
tff(f169,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Bool_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : (idiv_0(succ_0(X3),X4,X0) | ~diseqBool_0(X1,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X3,X2,X0) | ~lt_0(X1,X4,X0))),
  inference(flattening,[],[f168])).
tff(f168,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Bool_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : (idiv_0(succ_0(X3),X4,X0) | (~diseqBool_0(X1,true_0) | ~minus_0(X2,X4,X0) | ~idiv_0(X3,X2,X0) | ~lt_0(X1,X4,X0)))),
  inference(ennf_transformation,[],[f123])).
tff(f123,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Bool_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()'] : ((diseqBool_0(X1,true_0) & minus_0(X2,X4,X0) & idiv_0(X3,X2,X0) & lt_0(X1,X4,X0)) => idiv_0(succ_0(X3),X4,X0))),
  inference(rectify,[],[f62])).
tff(f62,axiom,(
  ! [X4 : 'Nat_0()',X2 : 'Bool_0()',X0 : 'Nat_0()',X1 : 'Nat_0()',X3 : 'Nat_0()'] : ((diseqBool_0(X2,true_0) & minus_0(X0,X3,X4) & idiv_0(X1,X0,X4) & lt_0(X2,X3,X4)) => idiv_0(succ_0(X1),X3,X4))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f401,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(zero_0))),succ_0(succ_0(succ_0(X0))))) )),
  inference(unit_resulting_resolution,[],[f399,f243])).
tff(f243,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (minus_0(X2,succ_0(X1),succ_0(X0)) | ~minus_0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f147])).
tff(f147,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X2,succ_0(X1),succ_0(X0)) | ~minus_0(X2,X1,X0))),
  inference(ennf_transformation,[],[f104])).
tff(f104,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (minus_0(X2,X1,X0) => minus_0(X2,succ_0(X1),succ_0(X0)))),
  inference(rectify,[],[f45])).
tff(f45,axiom,(
  ! [X1 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()'] : (minus_0(X0,X2,X1) => minus_0(X0,succ_0(X2),succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f399,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(zero_0)),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f397,f243])).
tff(f397,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(zero_0),succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f211,f243])).
tff(f211,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f47])).
tff(f47,axiom,(
  ! [X0 : 'Nat_0()'] : minus_0(zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f485,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP2(succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f188,f410,f273])).
tff(f410,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f406,f244])).
tff(f244,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Bool_0()'] : (lt_0(X1,succ_0(X0),succ_0(X2)) | ~lt_0(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f148])).
tff(f148,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Bool_0()',X2 : 'Nat_0()'] : (lt_0(X1,succ_0(X0),succ_0(X2)) | ~lt_0(X1,X0,X2))),
  inference(ennf_transformation,[],[f105])).
tff(f105,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Bool_0()',X2 : 'Nat_0()'] : (lt_0(X1,X0,X2) => lt_0(X1,succ_0(X0),succ_0(X2)))),
  inference(rectify,[],[f48])).
tff(f48,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Bool_0()',X2 : 'Nat_0()'] : (lt_0(X0,X1,X2) => lt_0(X0,succ_0(X1),succ_0(X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f406,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(succ_0(X0)),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f403,f244])).
tff(f403,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,succ_0(X0),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f212,f244])).
tff(f212,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(false_0,X0,zero_0)) )),
  inference(cnf_transformation,[],[f50])).
tff(f50,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(false_0,X0,zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f188,plain,(
  diseqBool_0(false_0,true_0)),
  inference(cnf_transformation,[],[f3])).
tff(f3,axiom,(
  diseqBool_0(false_0,true_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f341,plain,(
  ( ! [X0 : 'Nat_0()'] : (idiv_0(zero_0,zero_0,succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f221,f233])).
tff(f233,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~lt_0(true_0,X1,X0) | idiv_0(zero_0,X1,X0)) )),
  inference(cnf_transformation,[],[f137])).
tff(f137,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (idiv_0(zero_0,X1,X0) | ~lt_0(true_0,X1,X0))),
  inference(ennf_transformation,[],[f93])).
tff(f93,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (lt_0(true_0,X1,X0) => idiv_0(zero_0,X1,X0))),
  inference(rectify,[],[f61])).
tff(f61,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : (lt_0(true_0,X0,X1) => idiv_0(zero_0,X0,X1))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f221,plain,(
  ( ! [X0 : 'Nat_0()'] : (lt_0(true_0,zero_0,succ_0(X0))) )),
  inference(cnf_transformation,[],[f49])).
tff(f49,axiom,(
  ! [X0 : 'Nat_0()'] : lt_0(true_0,zero_0,succ_0(X0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f648,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(succ_0(succ_0(succ_0(zero_0))),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f644,f643])).
tff(f643,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (length_0(succ_0(X0),cons_0(X2,X1)) | ~length_0(X0,X1)) )),
  inference(resolution,[],[f637,f272])).
tff(f272,plain,(
  ( ! [X2 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (~sP1(X2,X1) | length_0(X2,cons_0(X3,X1))) )),
  inference(general_splitting,[],[f259,f271_D])).
tff(f271,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'list_0()'] : (~plus_0(X2,succ_0(zero_0),X0) | ~length_0(X0,X1) | sP1(X2,X1)) )),
  inference(cnf_transformation,[],[f271_D])).
tff(f271_D,plain,(
  ( ! [X1,X2] : (( ! [X0] : (~plus_0(X2,succ_0(zero_0),X0) | ~length_0(X0,X1)) ) <=> ~sP1(X2,X1)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP1])])).
tff(f259,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (length_0(X2,cons_0(X3,X1)) | ~length_0(X0,X1) | ~plus_0(X2,succ_0(zero_0),X0)) )),
  inference(cnf_transformation,[],[f167])).
tff(f167,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (length_0(X2,cons_0(X3,X1)) | ~length_0(X0,X1) | ~plus_0(X2,succ_0(zero_0),X0))),
  inference(flattening,[],[f166])).
tff(f166,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (length_0(X2,cons_0(X3,X1)) | (~length_0(X0,X1) | ~plus_0(X2,succ_0(zero_0),X0)))),
  inference(ennf_transformation,[],[f122])).
tff(f122,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : ((length_0(X0,X1) & plus_0(X2,succ_0(zero_0),X0)) => length_0(X2,cons_0(X3,X1)))),
  inference(rectify,[],[f59])).
tff(f59,axiom,(
  ! [X1 : 'Nat_0()',X3 : 'list_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : ((length_0(X1,X3) & plus_0(X0,succ_0(zero_0),X1)) => length_0(X0,cons_0(X2,X3)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f637,plain,(
  ( ! [X4 : 'list_0()',X3 : 'Nat_0()'] : (sP1(succ_0(X3),X4) | ~length_0(X3,X4)) )),
  inference(resolution,[],[f271,f369])).
tff(f369,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f215,f242])).
tff(f242,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (plus_0(succ_0(X1),succ_0(X0),X2) | ~plus_0(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f146])).
tff(f146,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(succ_0(X1),succ_0(X0),X2) | ~plus_0(X1,X0,X2))),
  inference(ennf_transformation,[],[f103])).
tff(f103,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X1,X0,X2) => plus_0(succ_0(X1),succ_0(X0),X2))),
  inference(rectify,[],[f41])).
tff(f41,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X0,X1,X2) => plus_0(succ_0(X0),succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f215,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f42])).
tff(f42,axiom,(
  ! [X0 : 'Nat_0()'] : plus_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f644,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (length_0(succ_0(succ_0(zero_0)),cons_0(X0,cons_0(X1,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f640,f272])).
tff(f640,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP1(succ_0(succ_0(zero_0)),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f369,f638,f271])).
tff(f638,plain,(
  ( ! [X0 : 'Nat_0()'] : (length_0(succ_0(zero_0),cons_0(X0,nil_0))) )),
  inference(unit_resulting_resolution,[],[f635,f272])).
tff(f635,plain,(
  sP1(succ_0(zero_0),nil_0)),
  inference(unit_resulting_resolution,[],[f189,f369,f271])).
tff(f189,plain,(
  length_0(zero_0,nil_0)),
  inference(cnf_transformation,[],[f60])).
tff(f60,axiom,(
  length_0(zero_0,nil_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f3897,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (splitAt_0(pair_1(cons_0(X0,nil_0),X1),succ_0(zero_0),cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f489,f1028,f258])).
tff(f258,plain,(
  ( ! [X2 : 'list_0()',X0 : 'list_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (splitAt_0(pair_1(X3,X0),X1,X2) | ~take_0(X3,X1,X2) | ~drop_0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f165])).
tff(f165,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()'] : (splitAt_0(pair_1(X3,X0),X1,X2) | ~take_0(X3,X1,X2) | ~drop_0(X0,X1,X2))),
  inference(flattening,[],[f164])).
tff(f164,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()'] : (splitAt_0(pair_1(X3,X0),X1,X2) | (~take_0(X3,X1,X2) | ~drop_0(X0,X1,X2)))),
  inference(ennf_transformation,[],[f121])).
tff(f121,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()'] : ((take_0(X3,X1,X2) & drop_0(X0,X1,X2)) => splitAt_0(pair_1(X3,X0),X1,X2))),
  inference(rectify,[],[f66])).
tff(f66,axiom,(
  ! [X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()',X0 : 'list_0()'] : ((take_0(X0,X2,X3) & drop_0(X1,X2,X3)) => splitAt_0(pair_1(X0,X1),X2,X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f1028,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (take_0(cons_0(X0,nil_0),succ_0(zero_0),cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f210,f254])).
tff(f254,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'list_0()'] : (take_0(cons_0(X0,X2),succ_0(X3),cons_0(X0,X1)) | ~take_0(X2,X3,X1)) )),
  inference(cnf_transformation,[],[f159])).
tff(f159,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (take_0(cons_0(X0,X2),succ_0(X3),cons_0(X0,X1)) | ~take_0(X2,X3,X1))),
  inference(ennf_transformation,[],[f116])).
tff(f116,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'Nat_0()'] : (take_0(X2,X3,X1) => take_0(cons_0(X0,X2),succ_0(X3),cons_0(X0,X1)))),
  inference(rectify,[],[f38])).
tff(f38,axiom,(
  ! [X1 : 'Nat_0()',X2 : 'list_0()',X0 : 'list_0()',X3 : 'Nat_0()'] : (take_0(X0,X3,X2) => take_0(cons_0(X1,X0),succ_0(X3),cons_0(X1,X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f210,plain,(
  ( ! [X0 : 'list_0()'] : (take_0(nil_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f40])).
tff(f40,axiom,(
  ! [X0 : 'list_0()'] : take_0(nil_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f489,plain,(
  ( ! [X0 : 'list_0()',X1 : 'Nat_0()'] : (drop_0(X0,succ_0(zero_0),cons_0(X1,X0))) )),
  inference(unit_resulting_resolution,[],[f216,f253])).
tff(f253,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X3 : 'list_0()',X1 : 'Nat_0()'] : (drop_0(X3,succ_0(X0),cons_0(X1,X2)) | ~drop_0(X3,X0,X2)) )),
  inference(cnf_transformation,[],[f158])).
tff(f158,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()'] : (drop_0(X3,succ_0(X0),cons_0(X1,X2)) | ~drop_0(X3,X0,X2))),
  inference(ennf_transformation,[],[f115])).
tff(f115,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'list_0()'] : (drop_0(X3,X0,X2) => drop_0(X3,succ_0(X0),cons_0(X1,X2)))),
  inference(rectify,[],[f63])).
tff(f63,axiom,(
  ! [X3 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'list_0()',X0 : 'list_0()'] : (drop_0(X0,X3,X2) => drop_0(X0,succ_0(X3),cons_0(X1,X2)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f216,plain,(
  ( ! [X0 : 'list_0()'] : (drop_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f65])).
tff(f65,axiom,(
  ! [X0 : 'list_0()'] : drop_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f853,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP10(cons_0(X0,nil_0),cons_0(X0,cons_0(zero_0,cons_0(X1,nil_0))),cons_0(zero_0,cons_0(X1,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f833,f846,f289])).
tff(f846,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (x_13(cons_0(X0,X1),cons_0(X0,nil_0),X1)) )),
  inference(unit_resulting_resolution,[],[f214,f252])).
tff(f252,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (x_13(cons_0(X2,X3),cons_0(X2,X0),X1) | ~x_13(X3,X0,X1)) )),
  inference(cnf_transformation,[],[f157])).
tff(f157,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (x_13(cons_0(X2,X3),cons_0(X2,X0),X1) | ~x_13(X3,X0,X1))),
  inference(ennf_transformation,[],[f114])).
tff(f114,plain,(
  ! [X0 : 'list_0()',X1 : 'list_0()',X2 : 'Nat_0()',X3 : 'list_0()'] : (x_13(X3,X0,X1) => x_13(cons_0(X2,X3),cons_0(X2,X0),X1))),
  inference(rectify,[],[f67])).
tff(f67,axiom,(
  ! [X2 : 'list_0()',X3 : 'list_0()',X1 : 'Nat_0()',X0 : 'list_0()'] : (x_13(X0,X2,X3) => x_13(cons_0(X1,X0),cons_0(X1,X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f214,plain,(
  ( ! [X0 : 'list_0()'] : (x_13(X0,nil_0,X0)) )),
  inference(cnf_transformation,[],[f68])).
tff(f68,axiom,(
  ! [X0 : 'list_0()'] : x_13(X0,nil_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f833,plain,(
  ( ! [X0 : 'Nat_0()'] : (stoogesort_1(cons_0(zero_0,cons_0(X0,nil_0)),cons_0(zero_0,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f827,f241])).
tff(f241,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'list_0()',X1 : 'Nat_0()'] : (stoogesort_1(X0,cons_0(X1,cons_0(X2,nil_0))) | ~sort_0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f145])).
tff(f145,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (stoogesort_1(X0,cons_0(X1,cons_0(X2,nil_0))) | ~sort_0(X0,X1,X2))),
  inference(ennf_transformation,[],[f102])).
tff(f102,plain,(
  ! [X0 : 'list_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (sort_0(X0,X1,X2) => stoogesort_1(X0,cons_0(X1,cons_0(X2,nil_0))))),
  inference(rectify,[],[f71])).
tff(f71,axiom,(
  ! [X0 : 'list_0()',X2 : 'Nat_0()',X1 : 'Nat_0()'] : (sort_0(X0,X2,X1) => stoogesort_1(X0,cons_0(X2,cons_0(X1,nil_0))))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f827,plain,(
  ( ! [X0 : 'Nat_0()'] : (sort_0(cons_0(zero_0,cons_0(X0,nil_0)),zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f213,f234])).
tff(f234,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sort_0(cons_0(X0,cons_0(X1,nil_0)),X0,X1) | ~leq_0(true_0,X0,X1)) )),
  inference(cnf_transformation,[],[f138])).
tff(f138,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sort_0(cons_0(X0,cons_0(X1,nil_0)),X0,X1) | ~leq_0(true_0,X0,X1))),
  inference(ennf_transformation,[],[f57])).
tff(f57,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (leq_0(true_0,X0,X1) => sort_0(cons_0(X0,cons_0(X1,nil_0)),X0,X1))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f213,plain,(
  ( ! [X0 : 'Nat_0()'] : (leq_0(true_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f53])).
tff(f53,axiom,(
  ! [X0 : 'Nat_0()'] : leq_0(true_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f7271,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~stoogesort_1(cons_0(succ_0(X0),cons_0(zero_0,cons_0(X1,nil_0))),X2)) )),
  inference(unit_resulting_resolution,[],[f7267,f270])).
tff(f270,plain,(
  ( ! [X2 : 'list_0()',X1 : 'list_0()'] : (~stoogesort_1(X2,X1) | ~sP0(X2)) )),
  inference(general_splitting,[],[f246,f269_D])).
tff(f269,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Bool_0()'] : (~ordered_0(X0,X2) | ~diseqBool_0(X0,true_0) | sP0(X2)) )),
  inference(cnf_transformation,[],[f269_D])).
tff(f269_D,plain,(
  ( ! [X2] : (( ! [X0] : (~ordered_0(X0,X2) | ~diseqBool_0(X0,true_0)) ) <=> ~sP0(X2)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP0])])).
tff(f246,plain,(
  ( ! [X2 : 'list_0()',X0 : 'Bool_0()',X1 : 'list_0()'] : (~diseqBool_0(X0,true_0) | ~stoogesort_1(X2,X1) | ~ordered_0(X0,X2)) )),
  inference(cnf_transformation,[],[f150])).
tff(f150,plain,(
  ! [X0 : 'Bool_0()',X1 : 'list_0()',X2 : 'list_0()'] : (~diseqBool_0(X0,true_0) | ~stoogesort_1(X2,X1) | ~ordered_0(X0,X2))),
  inference(ennf_transformation,[],[f108])).
tff(f108,plain,(
  ! [X0 : 'Bool_0()',X1 : 'list_0()',X2 : 'list_0()'] : ~(diseqBool_0(X0,true_0) & stoogesort_1(X2,X1) & ordered_0(X0,X2))),
  inference(true_and_false_elimination,[],[f107])).
tff(f107,plain,(
  ! [X0 : 'Bool_0()',X1 : 'list_0()',X2 : 'list_0()'] : ((diseqBool_0(X0,true_0) & stoogesort_1(X2,X1) & ordered_0(X0,X2)) => $false)),
  inference(rectify,[],[f87])).
tff(f87,axiom,(
  ! [X1 : 'Bool_0()',X2 : 'list_0()',X0 : 'list_0()'] : ((diseqBool_0(X1,true_0) & stoogesort_1(X0,X2) & ordered_0(X1,X0)) => $false)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f7267,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP0(cons_0(succ_0(X0),cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f188,f7184,f269])).
tff(f7184,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (ordered_0(false_0,cons_0(succ_0(X0),cons_0(zero_0,cons_0(X1,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f1336,f7163,f275])).
tff(f275,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'list_0()',X5 : 'Bool_0()',X3 : 'Bool_0()',X1 : 'Nat_0()'] : (~ordered_0(X5,cons_0(X4,X2)) | ordered_0(X3,cons_0(X1,cons_0(X4,X2))) | sP3(X5,X4,X3,X1)) )),
  inference(cnf_transformation,[],[f275_D])).
tff(f275_D,plain,(
  ( ! [X1,X3,X4,X5] : (( ! [X2] : (~ordered_0(X5,cons_0(X4,X2)) | ordered_0(X3,cons_0(X1,cons_0(X4,X2)))) ) <=> ~sP3(X5,X4,X3,X1)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP3])])).
tff(f7163,plain,(
  ( ! [X0 : 'Nat_0()'] : (ordered_0(true_0,cons_0(zero_0,cons_0(X0,nil_0)))) )),
  inference(unit_resulting_resolution,[],[f1345,f217,f275])).
tff(f217,plain,(
  ( ! [X0 : 'Nat_0()'] : (ordered_0(true_0,cons_0(X0,nil_0))) )),
  inference(cnf_transformation,[],[f55])).
tff(f55,axiom,(
  ! [X0 : 'Nat_0()'] : ordered_0(true_0,cons_0(X0,nil_0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f1345,plain,(
  ( ! [X0 : 'Nat_0()'] : (~sP3(true_0,X0,true_0,zero_0)) )),
  inference(unit_resulting_resolution,[],[f213,f204,f276])).
tff(f276,plain,(
  ( ! [X4 : 'Nat_0()',X0 : 'Bool_0()',X5 : 'Bool_0()',X3 : 'Bool_0()',X1 : 'Nat_0()'] : (~sP3(X5,X4,X3,X1) | ~and_0(X3,X0,X5) | ~leq_0(X0,X1,X4)) )),
  inference(general_splitting,[],[f261,f275_D])).
tff(f261,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'list_0()',X0 : 'Bool_0()',X5 : 'Bool_0()',X3 : 'Bool_0()',X1 : 'Nat_0()'] : (ordered_0(X3,cons_0(X1,cons_0(X4,X2))) | ~leq_0(X0,X1,X4) | ~ordered_0(X5,cons_0(X4,X2)) | ~and_0(X3,X0,X5)) )),
  inference(cnf_transformation,[],[f171])).
tff(f171,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Bool_0()',X4 : 'Nat_0()',X5 : 'Bool_0()'] : (ordered_0(X3,cons_0(X1,cons_0(X4,X2))) | ~leq_0(X0,X1,X4) | ~ordered_0(X5,cons_0(X4,X2)) | ~and_0(X3,X0,X5))),
  inference(flattening,[],[f170])).
tff(f170,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Bool_0()',X4 : 'Nat_0()',X5 : 'Bool_0()'] : (ordered_0(X3,cons_0(X1,cons_0(X4,X2))) | (~leq_0(X0,X1,X4) | ~ordered_0(X5,cons_0(X4,X2)) | ~and_0(X3,X0,X5)))),
  inference(ennf_transformation,[],[f124])).
tff(f124,plain,(
  ! [X0 : 'Bool_0()',X1 : 'Nat_0()',X2 : 'list_0()',X3 : 'Bool_0()',X4 : 'Nat_0()',X5 : 'Bool_0()'] : ((leq_0(X0,X1,X4) & ordered_0(X5,cons_0(X4,X2)) & and_0(X3,X0,X5)) => ordered_0(X3,cons_0(X1,cons_0(X4,X2))))),
  inference(rectify,[],[f54])).
tff(f54,axiom,(
  ! [X1 : 'Bool_0()',X5 : 'Nat_0()',X4 : 'list_0()',X0 : 'Bool_0()',X3 : 'Nat_0()',X2 : 'Bool_0()'] : ((leq_0(X1,X5,X3) & ordered_0(X2,cons_0(X3,X4)) & and_0(X0,X1,X2)) => ordered_0(X0,cons_0(X5,cons_0(X3,X4))))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f204,plain,(
  and_0(true_0,true_0,true_0)),
  inference(cnf_transformation,[],[f8])).
tff(f8,axiom,(
  and_0(true_0,true_0,true_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f1336,plain,(
  ( ! [X0 : 'Nat_0()'] : (~sP3(true_0,zero_0,false_0,succ_0(X0))) )),
  inference(unit_resulting_resolution,[],[f222,f202,f276])).
tff(f202,plain,(
  and_0(false_0,false_0,true_0)),
  inference(cnf_transformation,[],[f7])).
tff(f7,axiom,(
  and_0(false_0,false_0,true_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f222,plain,(
  ( ! [X0 : 'Nat_0()'] : (leq_0(false_0,succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f52])).
tff(f52,axiom,(
  ! [X0 : 'Nat_0()'] : leq_0(false_0,succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f33969,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (stoogesort_0(cons_0(X0,cons_0(X1,cons_0(X2,nil_0))),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f850,f11521,f312])).
tff(f312,plain,(
  ( ! [X4 : 'list_0()',X5 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (~sP21(X5,X4,X3) | ~sP20(X5,X4,X1) | stoogesort_0(X1,X3)) )),
  inference(general_splitting,[],[f310,f311_D])).
tff(f311,plain,(
  ( ! [X4 : 'list_0()',X0 : 'Nat_0()',X5 : 'list_0()',X3 : 'list_0()'] : (~splitAt_0(pair_1(X5,X4),X0,X3) | ~sP19(X3,X0) | sP21(X5,X4,X3)) )),
  inference(cnf_transformation,[],[f311_D])).
tff(f311_D,plain,(
  ( ! [X3,X4,X5] : (( ! [X0] : (~splitAt_0(pair_1(X5,X4),X0,X3) | ~sP19(X3,X0)) ) <=> ~sP21(X5,X4,X3)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP21])])).
tff(f310,plain,(
  ( ! [X4 : 'list_0()',X0 : 'Nat_0()',X5 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X1,X3) | ~splitAt_0(pair_1(X5,X4),X0,X3) | ~sP19(X3,X0) | ~sP20(X5,X4,X1)) )),
  inference(general_splitting,[],[f308,f309_D])).
tff(f309,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X5 : 'list_0()',X1 : 'list_0()'] : (~x_13(X1,X2,X4) | ~stoogesort_1(X2,X5) | sP20(X5,X4,X1)) )),
  inference(cnf_transformation,[],[f309_D])).
tff(f309_D,plain,(
  ( ! [X1,X4,X5] : (( ! [X2] : (~x_13(X1,X2,X4) | ~stoogesort_1(X2,X5)) ) <=> ~sP20(X5,X4,X1)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP20])])).
tff(f308,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'Nat_0()',X5 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X1,X3) | ~stoogesort_1(X2,X5) | ~x_13(X1,X2,X4) | ~splitAt_0(pair_1(X5,X4),X0,X3) | ~sP19(X3,X0)) )),
  inference(general_splitting,[],[f306,f307_D])).
tff(f307,plain,(
  ( ! [X0 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'list_0()'] : (~sP18(X0,X7) | ~length_0(X7,X3) | sP19(X3,X0)) )),
  inference(cnf_transformation,[],[f307_D])).
tff(f307_D,plain,(
  ( ! [X0,X3] : (( ! [X7] : (~sP18(X0,X7) | ~length_0(X7,X3)) ) <=> ~sP19(X3,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP19])])).
tff(f306,plain,(
  ( ! [X4 : 'list_0()',X2 : 'list_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X1,X3) | ~stoogesort_1(X2,X5) | ~x_13(X1,X2,X4) | ~length_0(X7,X3) | ~splitAt_0(pair_1(X5,X4),X0,X3) | ~sP18(X0,X7)) )),
  inference(general_splitting,[],[f268,f305_D])).
tff(f305,plain,(
  ( ! [X6 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()'] : (~idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~times_0(X6,succ_0(succ_0(zero_0)),X7) | sP18(X0,X7)) )),
  inference(cnf_transformation,[],[f305_D])).
tff(f305_D,plain,(
  ( ! [X7,X0] : (( ! [X6] : (~idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~times_0(X6,succ_0(succ_0(zero_0)),X7)) ) <=> ~sP18(X0,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP18])])).
tff(f268,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'list_0()',X2 : 'list_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'list_0()',X3 : 'list_0()',X1 : 'list_0()'] : (stoogesort_0(X1,X3) | ~stoogesort_1(X2,X5) | ~x_13(X1,X2,X4) | ~length_0(X7,X3) | ~times_0(X6,succ_0(succ_0(zero_0)),X7) | ~idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X5,X4),X0,X3)) )),
  inference(cnf_transformation,[],[f181])).
tff(f181,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (stoogesort_0(X1,X3) | ~stoogesort_1(X2,X5) | ~x_13(X1,X2,X4) | ~length_0(X7,X3) | ~times_0(X6,succ_0(succ_0(zero_0)),X7) | ~idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X5,X4),X0,X3))),
  inference(flattening,[],[f180])).
tff(f180,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (stoogesort_0(X1,X3) | (~stoogesort_1(X2,X5) | ~x_13(X1,X2,X4) | ~length_0(X7,X3) | ~times_0(X6,succ_0(succ_0(zero_0)),X7) | ~idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) | ~splitAt_0(pair_1(X5,X4),X0,X3)))),
  inference(ennf_transformation,[],[f135])).
tff(f135,plain,(
  ! [X0 : 'Nat_0()',X1 : 'list_0()',X2 : 'list_0()',X3 : 'list_0()',X4 : 'list_0()',X5 : 'list_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : ((stoogesort_1(X2,X5) & x_13(X1,X2,X4) & length_0(X7,X3) & times_0(X6,succ_0(succ_0(zero_0)),X7) & idiv_0(X0,succ_0(X6),succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X5,X4),X0,X3)) => stoogesort_0(X1,X3))),
  inference(rectify,[],[f69])).
tff(f69,axiom,(
  ! [X4 : 'Nat_0()',X0 : 'list_0()',X1 : 'list_0()',X7 : 'list_0()',X6 : 'list_0()',X5 : 'list_0()',X3 : 'Nat_0()',X2 : 'Nat_0()'] : ((stoogesort_1(X1,X5) & x_13(X0,X1,X6) & length_0(X2,X7) & times_0(X3,succ_0(succ_0(zero_0)),X2) & idiv_0(X4,succ_0(X3),succ_0(succ_0(succ_0(zero_0)))) & splitAt_0(pair_1(X5,X6),X4,X7)) => stoogesort_0(X0,X7))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f11521,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP21(cons_0(X0,nil_0),cons_0(X1,cons_0(X2,nil_0)),cons_0(X0,cons_0(X1,cons_0(X2,nil_0))))) )),
  inference(unit_resulting_resolution,[],[f3897,f11486,f311])).
tff(f11486,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (sP19(cons_0(X0,cons_0(X1,cons_0(X2,nil_0))),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f648,f11377,f307])).
tff(f11377,plain,(
  sP18(succ_0(zero_0),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f6639,f1265,f305])).
tff(f1265,plain,(
  times_0(succ_0(succ_0(succ_0(succ_0(succ_0(succ_0(zero_0)))))),succ_0(succ_0(zero_0)),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f383,f1035,f256])).
tff(f256,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~times_0(X1,X2,X3) | times_0(X0,succ_0(X2),X3) | ~plus_0(X0,X3,X1)) )),
  inference(cnf_transformation,[],[f162])).
tff(f162,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X0,succ_0(X2),X3) | ~times_0(X1,X2,X3) | ~plus_0(X0,X3,X1))),
  inference(flattening,[],[f161])).
tff(f161,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X0,succ_0(X2),X3) | (~times_0(X1,X2,X3) | ~plus_0(X0,X3,X1)))),
  inference(ennf_transformation,[],[f43])).
tff(f43,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : ((times_0(X1,X2,X3) & plus_0(X0,X3,X1)) => times_0(X0,succ_0(X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f1035,plain,(
  times_0(succ_0(succ_0(succ_0(zero_0))),succ_0(zero_0),succ_0(succ_0(succ_0(zero_0))))),
  inference(unit_resulting_resolution,[],[f383,f209,f256])).
tff(f209,plain,(
  ( ! [X0 : 'Nat_0()'] : (times_0(zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f44])).
tff(f44,axiom,(
  ! [X0 : 'Nat_0()'] : times_0(zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f383,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(succ_0(succ_0(X0))),succ_0(succ_0(succ_0(zero_0))),X0)) )),
  inference(unit_resulting_resolution,[],[f380,f242])).
tff(f380,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(succ_0(X0)),succ_0(succ_0(zero_0)),X0)) )),
  inference(unit_resulting_resolution,[],[f369,f242])).
tff(f6639,plain,(
  ( ! [X0 : 'Nat_0()'] : (idiv_0(succ_0(zero_0),succ_0(succ_0(succ_0(succ_0(X0)))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f341,f485,f402,f274])).
tff(f402,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(succ_0(X0)))),succ_0(succ_0(succ_0(zero_0))))) )),
  inference(unit_resulting_resolution,[],[f400,f243])).
tff(f400,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(succ_0(X0))),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f398,f243])).
tff(f398,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(succ_0(X0)),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f220,f243])).
tff(f220,plain,(
  ( ! [X0 : 'Nat_0()'] : (minus_0(zero_0,succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f46])).
tff(f46,axiom,(
  ! [X0 : 'Nat_0()'] : minus_0(zero_0,succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
tff(f850,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'list_0()'] : (sP20(cons_0(X0,nil_0),X1,cons_0(X0,X1))) )),
  inference(unit_resulting_resolution,[],[f223,f846,f309])).
tff(f223,plain,(
  ( ! [X0 : 'Nat_0()'] : (stoogesort_1(cons_0(X0,nil_0),cons_0(X0,nil_0))) )),
  inference(cnf_transformation,[],[f72])).
tff(f72,axiom,(
  ! [X0 : 'Nat_0()'] : stoogesort_1(cons_0(X0,nil_0),cons_0(X0,nil_0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/sort_nat_StoogeSort2Sorts.smt2',unknown)).
