8
20812
620
UNSAT
tff(type_def_5, type, 'Nat_0()': $tType).
tff(func_def_0, type, zero_0: 'Nat_0()').
tff(func_def_1, type, succ_0: ('Nat_0()') > 'Nat_0()').
tff(pred_def_1, type, diseqNat_0: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_2, type, p_1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_3, type, iszero_0: ('Nat_0()') > $o).
tff(pred_def_4, type, issucc_0: ('Nat_0()') > $o).
tff(pred_def_5, type, plus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_6, type, times_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_7, type, addacc_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_8, type, mulacc_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_9, type, sP0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_10, type, sP1: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_11, type, sP2: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_12, type, sP3: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_13, type, sP4: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_14, type, sP5: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_15, type, sP6: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_16, type, sP7: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_17, type, sP8: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_18, type, sP9: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_19, type, sP10: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_20, type, sP11: ('Nat_0()') > $o).
tff(pred_def_21, type, sP12: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_22, type, sP13: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_23, type, sP14: ('Nat_0()') > $o).
tff(pred_def_24, type, sP15: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_25, type, sP16: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_26, type, sP17: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_27, type, sP18: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_28, type, sP19: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_29, type, sP20: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_30, type, sP21: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_31, type, sP22: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_32, type, sP23: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_33, type, sP24: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_34, type, sP25: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_35, type, sP26: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_36, type, sP27: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_37, type, sP28: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_38, type, sP29: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_39, type, sP30: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_40, type, sP31: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_41, type, sP32: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_42, type, sP33: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_43, type, sP34: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_44, type, sP35: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_45, type, sP36: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_46, type, sP37: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_47, type, sP38: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_48, type, sP39: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_49, type, sP40: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_50, type, sP41: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_51, type, sP42: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(f27718,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f7019,f113,f118,f9867,f117,f652,f184])).
tff(f184,plain,(
  ( ! [X6 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~diseqNat_0(X6,zero_0) | ~addacc_0(X5,X3,X7,X8) | ~sP15(X6,X5) | ~sP18(X7,X6)) )),
  inference(general_splitting,[],[f182,f183_D])).
tff(f183,plain,(
  ( ! [X6 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~sP17(X1,X7,X6) | sP18(X7,X6)) )),
  inference(cnf_transformation,[],[f183_D])).
tff(f183_D,plain,(
  ( ! [X6,X7] : (( ! [X1] : (~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~sP17(X1,X7,X6)) ) <=> ~sP18(X7,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP18])])).
tff(f182,plain,(
  ( ! [X6 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~sP15(X6,X5) | ~sP17(X1,X7,X6)) )),
  inference(general_splitting,[],[f180,f181_D])).
tff(f181,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~sP16(X4,X1,X7) | sP17(X1,X7,X6)) )),
  inference(cnf_transformation,[],[f181_D])).
tff(f181_D,plain,(
  ( ! [X6,X7,X1] : (( ! [X4] : (~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~sP16(X4,X1,X7)) ) <=> ~sP17(X1,X7,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP17])])).
tff(f180,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~sP15(X6,X5) | ~sP16(X4,X1,X7)) )),
  inference(general_splitting,[],[f178,f179_D])).
tff(f179,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2) | sP16(X4,X1,X7)) )),
  inference(cnf_transformation,[],[f179_D])).
tff(f179_D,plain,(
  ( ! [X7,X1,X4] : (( ! [X2] : (~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2)) ) <=> ~sP16(X4,X1,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP16])])).
tff(f178,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~sP15(X6,X5)) )),
  inference(general_splitting,[],[f140,f177_D])).
tff(f177,plain,(
  ( ! [X6 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()'] : (~plus_0(X0,succ_0(zero_0),X5) | mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | sP15(X6,X5)) )),
  inference(cnf_transformation,[],[f177_D])).
tff(f177_D,plain,(
  ( ! [X5,X6] : (( ! [X0] : (~plus_0(X0,succ_0(zero_0),X5) | mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6))) ) <=> ~sP15(X6,X5)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP15])])).
tff(f140,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~plus_0(X0,succ_0(zero_0),X5)) )),
  inference(cnf_transformation,[],[f97])).
tff(f97,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : (mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~plus_0(X0,succ_0(zero_0),X5))),
  inference(flattening,[],[f96])).
tff(f96,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : (mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | (~diseqNat_0(X6,zero_0) | ~mulacc_0(X3,zero_0,zero_0,X6) | ~mulacc_0(X4,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X1,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X4,X1,X2) | ~addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X5,X3,X7,X8) | ~plus_0(X0,succ_0(zero_0),X5)))),
  inference(ennf_transformation,[],[f68])).
tff(f68,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : ((diseqNat_0(X6,zero_0) & mulacc_0(X3,zero_0,zero_0,X6) & mulacc_0(X4,succ_0(zero_0),zero_0,X6) & mulacc_0(X1,zero_0,succ_0(zero_0),X6) & mulacc_0(X2,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X7,X4,X1,X2) & addacc_0(X8,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) & addacc_0(X5,X3,X7,X8) & plus_0(X0,succ_0(zero_0),X5)) => mulacc_0(X0,succ_0(zero_0),succ_0(zero_0),succ_0(X6)))),
  inference(rectify,[],[f18])).
tff(f18,axiom,(
  ! [X7 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()',X6 : 'Nat_0()',X8 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()'] : ((diseqNat_0(X8,zero_0) & mulacc_0(X0,zero_0,zero_0,X8) & mulacc_0(X1,succ_0(zero_0),zero_0,X8) & mulacc_0(X2,zero_0,succ_0(zero_0),X8) & mulacc_0(X3,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X4,X1,X2,X3) & addacc_0(X5,succ_0(zero_0),succ_0(zero_0),succ_0(X8)) & addacc_0(X6,X0,X4,X5) & plus_0(X7,succ_0(zero_0),X6)) => mulacc_0(X7,succ_0(zero_0),succ_0(zero_0),succ_0(X8)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f652,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(succ_0(X0)),succ_0(zero_0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f645,f130])).
tff(f130,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~addacc_0(X0,X1,succ_0(X2),X3) | addacc_0(X0,succ_0(X1),X2,X3)) )),
  inference(cnf_transformation,[],[f84])).
tff(f84,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (addacc_0(X0,succ_0(X1),X2,X3) | ~addacc_0(X0,X1,succ_0(X2),X3))),
  inference(ennf_transformation,[],[f11])).
tff(f11,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (addacc_0(X0,X1,succ_0(X2),X3) => addacc_0(X0,succ_0(X1),X2,X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f645,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(succ_0(X0)),zero_0,succ_0(succ_0(zero_0)),X0)) )),
  inference(unit_resulting_resolution,[],[f643,f129])).
tff(f129,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~addacc_0(X2,zero_0,X1,succ_0(X0)) | addacc_0(X2,zero_0,succ_0(X1),X0)) )),
  inference(cnf_transformation,[],[f83])).
tff(f83,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (addacc_0(X2,zero_0,succ_0(X1),X0) | ~addacc_0(X2,zero_0,X1,succ_0(X0)))),
  inference(ennf_transformation,[],[f51])).
tff(f51,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (addacc_0(X2,zero_0,X1,succ_0(X0)) => addacc_0(X2,zero_0,succ_0(X1),X0))),
  inference(rectify,[],[f12])).
tff(f12,axiom,(
  ! [X2 : 'Nat_0()',X1 : 'Nat_0()',X0 : 'Nat_0()'] : (addacc_0(X0,zero_0,X1,succ_0(X2)) => addacc_0(X0,zero_0,succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f643,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(X0),zero_0,succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f117,f129])).
tff(f117,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(X0,zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f13])).
tff(f13,axiom,(
  ! [X0 : 'Nat_0()'] : addacc_0(X0,zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f9867,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP15(succ_0(zero_0),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f277,f1168,f177])).
tff(f1168,plain,(
  ( ! [X0 : 'Nat_0()'] : (~mulacc_0(succ_0(succ_0(succ_0(X0))),succ_0(zero_0),succ_0(zero_0),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f803,f960,f148])).
tff(f148,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X0,X4,X2,X1) | ~times_0(X5,X4,X2) | ~sP0(X5,X1,X0)) )),
  inference(general_splitting,[],[f134,f147_D])).
tff(f147,plain,(
  ( ! [X0 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~times_0(X3,X5,X1) | ~diseqNat_0(X0,X3) | sP0(X5,X1,X0)) )),
  inference(cnf_transformation,[],[f147_D])).
tff(f147_D,plain,(
  ( ! [X0,X1,X5] : (( ! [X3] : (~times_0(X3,X5,X1) | ~diseqNat_0(X0,X3)) ) <=> ~sP0(X5,X1,X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP0])])).
tff(f134,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X0,X3) | ~mulacc_0(X0,X4,X2,X1) | ~times_0(X5,X4,X2) | ~times_0(X3,X5,X1)) )),
  inference(cnf_transformation,[],[f89])).
tff(f89,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()'] : (~diseqNat_0(X0,X3) | ~mulacc_0(X0,X4,X2,X1) | ~times_0(X5,X4,X2) | ~times_0(X3,X5,X1))),
  inference(ennf_transformation,[],[f58])).
tff(f58,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()'] : ~(diseqNat_0(X0,X3) & mulacc_0(X0,X4,X2,X1) & times_0(X5,X4,X2) & times_0(X3,X5,X1))),
  inference(true_and_false_elimination,[],[f57])).
tff(f57,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()'] : ((diseqNat_0(X0,X3) & mulacc_0(X0,X4,X2,X1) & times_0(X5,X4,X2) & times_0(X3,X5,X1)) => $false)),
  inference(rectify,[],[f25])).
tff(f25,axiom,(
  ! [X0 : 'Nat_0()',X5 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : ((diseqNat_0(X0,X2) & mulacc_0(X0,X3,X4,X5) & times_0(X1,X3,X4) & times_0(X2,X1,X5)) => $false)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f960,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP0(succ_0(zero_0),succ_0(succ_0(zero_0)),succ_0(succ_0(succ_0(X0))))) )),
  inference(unit_resulting_resolution,[],[f254,f804,f147])).
tff(f804,plain,(
  times_0(succ_0(succ_0(zero_0)),succ_0(zero_0),succ_0(succ_0(zero_0)))),
  inference(unit_resulting_resolution,[],[f290,f115,f132])).
tff(f132,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~times_0(X2,X0,X3) | times_0(X1,succ_0(X0),X3) | ~plus_0(X1,X3,X2)) )),
  inference(cnf_transformation,[],[f87])).
tff(f87,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X1,succ_0(X0),X3) | ~times_0(X2,X0,X3) | ~plus_0(X1,X3,X2))),
  inference(flattening,[],[f86])).
tff(f86,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (times_0(X1,succ_0(X0),X3) | (~times_0(X2,X0,X3) | ~plus_0(X1,X3,X2)))),
  inference(ennf_transformation,[],[f54])).
tff(f54,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : ((times_0(X2,X0,X3) & plus_0(X1,X3,X2)) => times_0(X1,succ_0(X0),X3))),
  inference(rectify,[],[f9])).
tff(f9,axiom,(
  ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()',X3 : 'Nat_0()'] : ((times_0(X1,X2,X3) & plus_0(X0,X3,X1)) => times_0(X0,succ_0(X2),X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f115,plain,(
  ( ! [X0 : 'Nat_0()'] : (times_0(zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f10])).
tff(f10,axiom,(
  ! [X0 : 'Nat_0()'] : times_0(zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f290,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(succ_0(X0)),succ_0(succ_0(zero_0)),X0)) )),
  inference(unit_resulting_resolution,[],[f277,f128])).
tff(f128,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (plus_0(succ_0(X1),succ_0(X0),X2) | ~plus_0(X1,X0,X2)) )),
  inference(cnf_transformation,[],[f82])).
tff(f82,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(succ_0(X1),succ_0(X0),X2) | ~plus_0(X1,X0,X2))),
  inference(ennf_transformation,[],[f50])).
tff(f50,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X1,X0,X2) => plus_0(succ_0(X1),succ_0(X0),X2))),
  inference(rectify,[],[f7])).
tff(f7,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X0,X1,X2) => plus_0(succ_0(X0),succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f254,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(succ_0(X0))),succ_0(succ_0(zero_0)))) )),
  inference(unit_resulting_resolution,[],[f248,f121])).
tff(f121,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),succ_0(X1)) | ~diseqNat_0(X0,X1)) )),
  inference(cnf_transformation,[],[f75])).
tff(f75,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),succ_0(X1)) | ~diseqNat_0(X0,X1))),
  inference(ennf_transformation,[],[f6])).
tff(f6,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(X0,X1) => diseqNat_0(succ_0(X0),succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f248,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(X0)),succ_0(zero_0))) )),
  inference(unit_resulting_resolution,[],[f113,f121])).
tff(f803,plain,(
  times_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0))),
  inference(unit_resulting_resolution,[],[f277,f115,f132])).
tff(f277,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f116,f128])).
tff(f116,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f8])).
tff(f8,axiom,(
  ! [X0 : 'Nat_0()'] : plus_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f118,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(zero_0,zero_0,X0,X1)) )),
  inference(cnf_transformation,[],[f24])).
tff(f24,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : mulacc_0(zero_0,zero_0,X0,X1)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f113,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f5])).
tff(f5,axiom,(
  ! [X0 : 'Nat_0()'] : diseqNat_0(succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
tff(f7019,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP18(zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f118,f7016,f183])).
tff(f7016,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP17(zero_0,zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f119,f7003,f181])).
tff(f7003,plain,(
  sP16(zero_0,zero_0,zero_0)),
  inference(unit_resulting_resolution,[],[f117,f118,f179])).
tff(f119,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(zero_0,succ_0(X1),zero_0,X0)) )),
  inference(cnf_transformation,[],[f38])).
tff(f38,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : mulacc_0(zero_0,succ_0(X1),zero_0,X0)),
  inference(rectify,[],[f23])).
tff(f23,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : mulacc_0(zero_0,succ_0(X0),zero_0,X1)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_spec.smt2',unknown)).
