6
17756
330
UNSAT
tff(type_def_5, type, 'Nat_0()': $tType).
tff(func_def_0, type, zero_0: 'Nat_0()').
tff(func_def_1, type, succ_0: ('Nat_0()') > 'Nat_0()').
tff(pred_def_1, type, diseqNat_0: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_2, type, p_1: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_3, type, iszero_0: ('Nat_0()') > $o).
tff(pred_def_4, type, issucc_0: ('Nat_0()') > $o).
tff(pred_def_5, type, plus_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_6, type, addacc_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_7, type, mulacc_0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_8, type, sP0: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_9, type, sP1: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_10, type, sP2: ('Nat_0()') > $o).
tff(pred_def_11, type, sP3: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_12, type, sP4: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_13, type, sP5: ('Nat_0()') > $o).
tff(pred_def_14, type, sP6: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_15, type, sP7: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_16, type, sP8: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_17, type, sP9: ('Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_18, type, sP10: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_19, type, sP11: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_20, type, sP12: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_21, type, sP13: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_22, type, sP14: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_23, type, sP15: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_24, type, sP16: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_25, type, sP17: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_26, type, sP18: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_27, type, sP19: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_28, type, sP20: ('Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_29, type, sP21: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_30, type, sP22: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_31, type, sP23: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_32, type, sP24: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_33, type, sP25: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_34, type, sP26: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_35, type, sP27: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_36, type, sP28: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_37, type, sP29: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_38, type, sP30: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_39, type, sP31: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_40, type, sP32: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_41, type, sP33: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_42, type, sP34: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_43, type, sP35: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(pred_def_44, type, sP36: ('Nat_0()' * 'Nat_0()' * 'Nat_0()' * 'Nat_0()') > $o).
tff(f20197,plain,(
  $false),
  inference(unit_resulting_resolution,[],[f8360,f81,f82,f82,f83,f15060,f113])).
tff(f113,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~sP5(X7)) )),
  inference(general_splitting,[],[f111,f112_D])).
tff(f112,plain,(
  ( ! [X7 : 'Nat_0()',X3 : 'Nat_0()'] : (~mulacc_0(X3,zero_0,zero_0,zero_0) | ~sP4(X3,X7) | sP5(X7)) )),
  inference(cnf_transformation,[],[f112_D])).
tff(f112_D,plain,(
  ( ! [X7] : (( ! [X3] : (~mulacc_0(X3,zero_0,zero_0,zero_0) | ~sP4(X3,X7)) ) <=> ~sP5(X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP5])])).
tff(f111,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~sP4(X3,X7)) )),
  inference(general_splitting,[],[f109,f110_D])).
tff(f110,plain,(
  ( ! [X0 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()'] : (~sP3(X3,X0,X7) | ~sP2(X0) | sP4(X3,X7)) )),
  inference(cnf_transformation,[],[f110_D])).
tff(f110_D,plain,(
  ( ! [X7,X3] : (( ! [X0] : (~sP3(X3,X0,X7) | ~sP2(X0)) ) <=> ~sP4(X3,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP4])])).
tff(f109,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~sP2(X0) | ~sP3(X3,X0,X7)) )),
  inference(general_splitting,[],[f107,f108_D])).
tff(f108,plain,(
  ( ! [X6 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()'] : (~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6) | sP3(X3,X0,X7)) )),
  inference(cnf_transformation,[],[f108_D])).
tff(f108_D,plain,(
  ( ! [X7,X0,X3] : (( ! [X6] : (~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6)) ) <=> ~sP3(X3,X0,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP3])])).
tff(f107,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6) | ~sP2(X0)) )),
  inference(general_splitting,[],[f93,f106_D])).
tff(f106,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~plus_0(X1,succ_0(zero_0),X0) | sP2(X0)) )),
  inference(cnf_transformation,[],[f106_D])).
tff(f106_D,plain,(
  ( ! [X0] : (( ! [X1] : ~plus_0(X1,succ_0(zero_0),X0) ) <=> ~sP2(X0)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP2])])).
tff(f93,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6) | ~plus_0(X1,succ_0(zero_0),X0)) )),
  inference(cnf_transformation,[],[f59])).
tff(f59,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6) | ~plus_0(X1,succ_0(zero_0),X0))),
  inference(flattening,[],[f58])).
tff(f58,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : (mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | (~mulacc_0(X3,zero_0,zero_0,zero_0) | ~mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) | ~mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) | ~mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X7,X2,X5,X4) | ~addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) | ~addacc_0(X0,X3,X7,X6) | ~plus_0(X1,succ_0(zero_0),X0)))),
  inference(ennf_transformation,[],[f40])).
tff(f40,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()'] : ((mulacc_0(X3,zero_0,zero_0,zero_0) & mulacc_0(X2,succ_0(zero_0),zero_0,zero_0) & mulacc_0(X5,zero_0,succ_0(zero_0),zero_0) & mulacc_0(X4,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X7,X2,X5,X4) & addacc_0(X6,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) & addacc_0(X0,X3,X7,X6) & plus_0(X1,succ_0(zero_0),X0)) => mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)))),
  inference(rectify,[],[f12])).
tff(f12,axiom,(
  ! [X6 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X2 : 'Nat_0()',X5 : 'Nat_0()',X4 : 'Nat_0()'] : ((mulacc_0(X0,zero_0,zero_0,zero_0) & mulacc_0(X1,succ_0(zero_0),zero_0,zero_0) & mulacc_0(X2,zero_0,succ_0(zero_0),zero_0) & mulacc_0(X3,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X4,X1,X2,X3) & addacc_0(X5,succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)) & addacc_0(X6,X0,X4,X5) & plus_0(X7,succ_0(zero_0),X6)) => mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f15060,plain,(
  ~mulacc_0(succ_0(zero_0),succ_0(zero_0),succ_0(zero_0),succ_0(zero_0))),
  inference(unit_resulting_resolution,[],[f14900,f15059,f4515])).
tff(f4515,plain,(
  ( ! [X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X4,X5,X1,X2) | ~mulacc_0(X0,X1,X2,X3) | sP12(X3,X4,X5,X0)) )),
  inference(resolution,[],[f124,f126])).
tff(f126,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()'] : (~sP11(X2,X1,X0,X7,X6) | sP12(X2,X0,X7,X6)) )),
  inference(cnf_transformation,[],[f126_D])).
tff(f126_D,plain,(
  ( ! [X6,X7,X0,X2] : (( ! [X1] : ~sP11(X2,X1,X0,X7,X6) ) <=> ~sP12(X2,X0,X7,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP12])])).
tff(f124,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()'] : (sP11(X2,X1,X0,X7,X6) | ~mulacc_0(X6,X1,X4,X2) | ~mulacc_0(X0,X7,X1,X4)) )),
  inference(cnf_transformation,[],[f124_D])).
tff(f124_D,plain,(
  ( ! [X6,X7,X0,X1,X2] : (( ! [X4] : (~mulacc_0(X6,X1,X4,X2) | ~mulacc_0(X0,X7,X1,X4)) ) <=> ~sP11(X2,X1,X0,X7,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP11])])).
tff(f15059,plain,(
  ( ! [X0 : 'Nat_0()'] : (~sP12(succ_0(succ_0(zero_0)),succ_0(zero_0),succ_0(zero_0),succ_0(succ_0(succ_0(X0))))) )),
  inference(unit_resulting_resolution,[],[f14900,f14938,f127])).
tff(f127,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()'] : (~sP12(X2,X0,X7,X6) | ~sP10(X2,X0,X8,X3) | ~mulacc_0(X3,X8,X7,X6)) )),
  inference(general_splitting,[],[f125,f126_D])).
tff(f125,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X3,X8,X7,X6) | ~sP10(X2,X0,X8,X3) | ~sP11(X2,X1,X0,X7,X6)) )),
  inference(general_splitting,[],[f123,f124_D])).
tff(f123,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X0,X7,X1,X4) | ~mulacc_0(X6,X1,X4,X2) | ~mulacc_0(X3,X8,X7,X6) | ~sP10(X2,X0,X8,X3)) )),
  inference(general_splitting,[],[f95,f122_D])).
tff(f122,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()'] : (~mulacc_0(X5,X8,X0,X2) | ~diseqNat_0(X5,X3) | sP10(X2,X0,X8,X3)) )),
  inference(cnf_transformation,[],[f122_D])).
tff(f122_D,plain,(
  ( ! [X3,X8,X0,X2] : (( ! [X5] : (~mulacc_0(X5,X8,X0,X2) | ~diseqNat_0(X5,X3)) ) <=> ~sP10(X2,X0,X8,X3)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP10])])).
tff(f95,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X5,X3) | ~mulacc_0(X0,X7,X1,X4) | ~mulacc_0(X5,X8,X0,X2) | ~mulacc_0(X6,X1,X4,X2) | ~mulacc_0(X3,X8,X7,X6)) )),
  inference(cnf_transformation,[],[f62])).
tff(f62,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : (~diseqNat_0(X5,X3) | ~mulacc_0(X0,X7,X1,X4) | ~mulacc_0(X5,X8,X0,X2) | ~mulacc_0(X6,X1,X4,X2) | ~mulacc_0(X3,X8,X7,X6))),
  inference(ennf_transformation,[],[f43])).
tff(f43,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : ~(diseqNat_0(X5,X3) & mulacc_0(X0,X7,X1,X4) & mulacc_0(X5,X8,X0,X2) & mulacc_0(X6,X1,X4,X2) & mulacc_0(X3,X8,X7,X6))),
  inference(true_and_false_elimination,[],[f42])).
tff(f42,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : ((diseqNat_0(X5,X3) & mulacc_0(X0,X7,X1,X4) & mulacc_0(X5,X8,X0,X2) & mulacc_0(X6,X1,X4,X2) & mulacc_0(X3,X8,X7,X6)) => $false)),
  inference(rectify,[],[f23])).
tff(f23,axiom,(
  ! [X0 : 'Nat_0()',X6 : 'Nat_0()',X8 : 'Nat_0()',X3 : 'Nat_0()',X7 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X5 : 'Nat_0()',X4 : 'Nat_0()'] : ((diseqNat_0(X1,X3) & mulacc_0(X0,X5,X6,X7) & mulacc_0(X1,X4,X0,X8) & mulacc_0(X2,X6,X7,X8) & mulacc_0(X3,X4,X5,X2)) => $false)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f14938,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP10(succ_0(succ_0(zero_0)),succ_0(zero_0),succ_0(zero_0),succ_0(succ_0(succ_0(succ_0(succ_0(succ_0(X0)))))))) )),
  inference(unit_resulting_resolution,[],[f506,f14900,f122])).
tff(f506,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(succ_0(succ_0(succ_0(zero_0))))),succ_0(succ_0(succ_0(succ_0(succ_0(succ_0(X0)))))))) )),
  inference(unit_resulting_resolution,[],[f233,f85])).
tff(f85,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),succ_0(X1)) | ~diseqNat_0(X0,X1)) )),
  inference(cnf_transformation,[],[f50])).
tff(f50,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),succ_0(X1)) | ~diseqNat_0(X0,X1))),
  inference(ennf_transformation,[],[f6])).
tff(f6,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (diseqNat_0(X0,X1) => diseqNat_0(succ_0(X0),succ_0(X1)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f233,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(succ_0(succ_0(zero_0)))),succ_0(succ_0(succ_0(succ_0(succ_0(X0))))))) )),
  inference(unit_resulting_resolution,[],[f194,f85])).
tff(f194,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(succ_0(zero_0))),succ_0(succ_0(succ_0(succ_0(X0)))))) )),
  inference(unit_resulting_resolution,[],[f188,f85])).
tff(f188,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(succ_0(zero_0)),succ_0(succ_0(succ_0(X0))))) )),
  inference(unit_resulting_resolution,[],[f184,f85])).
tff(f184,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(zero_0),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f77,f85])).
tff(f77,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(zero_0,succ_0(X0))) )),
  inference(cnf_transformation,[],[f4])).
tff(f4,axiom,(
  ! [X0 : 'Nat_0()'] : diseqNat_0(zero_0,succ_0(X0))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f14900,plain,(
  ( ! [X0 : 'Nat_0()'] : (mulacc_0(succ_0(succ_0(succ_0(succ_0(succ_0(X0))))),succ_0(zero_0),succ_0(zero_0),succ_0(succ_0(X0)))) )),
  inference(unit_resulting_resolution,[],[f200,f14896,f114])).
tff(f114,plain,(
  ( ! [X6 : 'Nat_0()',X7 : 'Nat_0()',X3 : 'Nat_0()'] : (~plus_0(X3,succ_0(zero_0),X7) | mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | sP6(X6,X7)) )),
  inference(cnf_transformation,[],[f114_D])).
tff(f114_D,plain,(
  ( ! [X7,X6] : (( ! [X3] : (~plus_0(X3,succ_0(zero_0),X7) | mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6))) ) <=> ~sP6(X6,X7)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP6])])).
tff(f14896,plain,(
  ( ! [X0 : 'Nat_0()'] : (~sP6(succ_0(X0),succ_0(succ_0(succ_0(succ_0(X0)))))) )),
  inference(unit_resulting_resolution,[],[f3297,f78,f82,f81,f488,f121])).
tff(f121,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X0 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~diseqNat_0(X6,zero_0) | ~addacc_0(X7,X0,X5,X4) | ~sP6(X6,X7) | ~sP9(X5,X6)) )),
  inference(general_splitting,[],[f119,f120_D])).
tff(f120,plain,(
  ( ! [X6 : 'Nat_0()',X8 : 'Nat_0()',X5 : 'Nat_0()'] : (~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~sP8(X8,X5,X6) | sP9(X5,X6)) )),
  inference(cnf_transformation,[],[f120_D])).
tff(f120_D,plain,(
  ( ! [X6,X5] : (( ! [X8] : (~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~sP8(X8,X5,X6)) ) <=> ~sP9(X5,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP9])])).
tff(f119,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~sP6(X6,X7) | ~sP8(X8,X5,X6)) )),
  inference(general_splitting,[],[f117,f118_D])).
tff(f118,plain,(
  ( ! [X6 : 'Nat_0()',X2 : 'Nat_0()',X8 : 'Nat_0()',X5 : 'Nat_0()'] : (~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~sP7(X8,X5,X2) | sP8(X8,X5,X6)) )),
  inference(cnf_transformation,[],[f118_D])).
tff(f118_D,plain,(
  ( ! [X6,X5,X8] : (( ! [X2] : (~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~sP7(X8,X5,X2)) ) <=> ~sP8(X8,X5,X6)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP8])])).
tff(f117,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~sP6(X6,X7) | ~sP7(X8,X5,X2)) )),
  inference(general_splitting,[],[f115,f116_D])).
tff(f116,plain,(
  ( ! [X2 : 'Nat_0()',X8 : 'Nat_0()',X5 : 'Nat_0()',X1 : 'Nat_0()'] : (~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1) | sP7(X8,X5,X2)) )),
  inference(cnf_transformation,[],[f116_D])).
tff(f116_D,plain,(
  ( ! [X2,X5,X8] : (( ! [X1] : (~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1)) ) <=> ~sP7(X8,X5,X2)) )),
  introduced(general_splitting_component_introduction,[new_symbols(naming,[sP7])])).
tff(f115,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X1 : 'Nat_0()'] : (~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~sP6(X6,X7)) )),
  inference(general_splitting,[],[f94,f114_D])).
tff(f94,plain,(
  ( ! [X6 : 'Nat_0()',X4 : 'Nat_0()',X2 : 'Nat_0()',X0 : 'Nat_0()',X8 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~plus_0(X3,succ_0(zero_0),X7)) )),
  inference(cnf_transformation,[],[f61])).
tff(f61,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : (mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~plus_0(X3,succ_0(zero_0),X7))),
  inference(flattening,[],[f60])).
tff(f60,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : (mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | (~diseqNat_0(X6,zero_0) | ~mulacc_0(X0,zero_0,zero_0,X6) | ~mulacc_0(X8,succ_0(zero_0),zero_0,X6) | ~mulacc_0(X2,zero_0,succ_0(zero_0),X6) | ~mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) | ~addacc_0(X5,X8,X2,X1) | ~addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) | ~addacc_0(X7,X0,X5,X4) | ~plus_0(X3,succ_0(zero_0),X7)))),
  inference(ennf_transformation,[],[f41])).
tff(f41,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()',X4 : 'Nat_0()',X5 : 'Nat_0()',X6 : 'Nat_0()',X7 : 'Nat_0()',X8 : 'Nat_0()'] : ((diseqNat_0(X6,zero_0) & mulacc_0(X0,zero_0,zero_0,X6) & mulacc_0(X8,succ_0(zero_0),zero_0,X6) & mulacc_0(X2,zero_0,succ_0(zero_0),X6) & mulacc_0(X1,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X5,X8,X2,X1) & addacc_0(X4,succ_0(zero_0),succ_0(zero_0),succ_0(X6)) & addacc_0(X7,X0,X5,X4) & plus_0(X3,succ_0(zero_0),X7)) => mulacc_0(X3,succ_0(zero_0),succ_0(zero_0),succ_0(X6)))),
  inference(rectify,[],[f16])).
tff(f16,axiom,(
  ! [X0 : 'Nat_0()',X3 : 'Nat_0()',X2 : 'Nat_0()',X7 : 'Nat_0()',X5 : 'Nat_0()',X4 : 'Nat_0()',X8 : 'Nat_0()',X6 : 'Nat_0()',X1 : 'Nat_0()'] : ((diseqNat_0(X8,zero_0) & mulacc_0(X0,zero_0,zero_0,X8) & mulacc_0(X1,succ_0(zero_0),zero_0,X8) & mulacc_0(X2,zero_0,succ_0(zero_0),X8) & mulacc_0(X3,zero_0,zero_0,succ_0(zero_0)) & addacc_0(X4,X1,X2,X3) & addacc_0(X5,succ_0(zero_0),succ_0(zero_0),succ_0(X8)) & addacc_0(X6,X0,X4,X5) & plus_0(X7,succ_0(zero_0),X6)) => mulacc_0(X7,succ_0(zero_0),succ_0(zero_0),succ_0(X8)))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f488,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(succ_0(X0)),succ_0(zero_0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f481,f90])).
tff(f90,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X3 : 'Nat_0()',X1 : 'Nat_0()'] : (~addacc_0(X1,X3,succ_0(X2),X0) | addacc_0(X1,succ_0(X3),X2,X0)) )),
  inference(cnf_transformation,[],[f55])).
tff(f55,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (addacc_0(X1,succ_0(X3),X2,X0) | ~addacc_0(X1,X3,succ_0(X2),X0))),
  inference(ennf_transformation,[],[f35])).
tff(f35,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()',X3 : 'Nat_0()'] : (addacc_0(X1,X3,succ_0(X2),X0) => addacc_0(X1,succ_0(X3),X2,X0))),
  inference(rectify,[],[f9])).
tff(f9,axiom,(
  ! [X3 : 'Nat_0()',X0 : 'Nat_0()',X2 : 'Nat_0()',X1 : 'Nat_0()'] : (addacc_0(X0,X1,succ_0(X2),X3) => addacc_0(X0,succ_0(X1),X2,X3))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f481,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(succ_0(X0)),zero_0,succ_0(succ_0(zero_0)),X0)) )),
  inference(unit_resulting_resolution,[],[f479,f89])).
tff(f89,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (~addacc_0(X1,zero_0,X2,succ_0(X0)) | addacc_0(X1,zero_0,succ_0(X2),X0)) )),
  inference(cnf_transformation,[],[f54])).
tff(f54,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (addacc_0(X1,zero_0,succ_0(X2),X0) | ~addacc_0(X1,zero_0,X2,succ_0(X0)))),
  inference(ennf_transformation,[],[f34])).
tff(f34,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (addacc_0(X1,zero_0,X2,succ_0(X0)) => addacc_0(X1,zero_0,succ_0(X2),X0))),
  inference(rectify,[],[f10])).
tff(f10,axiom,(
  ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (addacc_0(X0,zero_0,X1,succ_0(X2)) => addacc_0(X0,zero_0,succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f479,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(succ_0(X0),zero_0,succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f81,f89])).
tff(f78,plain,(
  ( ! [X0 : 'Nat_0()'] : (diseqNat_0(succ_0(X0),zero_0)) )),
  inference(cnf_transformation,[],[f5])).
tff(f5,axiom,(
  ! [X0 : 'Nat_0()'] : diseqNat_0(succ_0(X0),zero_0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f3297,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP9(zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f83,f3294,f120])).
tff(f3294,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP8(zero_0,zero_0,X0)) )),
  inference(unit_resulting_resolution,[],[f82,f3277,f118])).
tff(f3277,plain,(
  sP7(zero_0,zero_0,zero_0)),
  inference(unit_resulting_resolution,[],[f81,f82,f116])).
tff(f200,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(zero_0),X0)) )),
  inference(unit_resulting_resolution,[],[f80,f88])).
tff(f88,plain,(
  ( ! [X2 : 'Nat_0()',X0 : 'Nat_0()',X1 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(X1),X2) | ~plus_0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f53])).
tff(f53,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(succ_0(X0),succ_0(X1),X2) | ~plus_0(X0,X1,X2))),
  inference(ennf_transformation,[],[f7])).
tff(f7,axiom,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()',X2 : 'Nat_0()'] : (plus_0(X0,X1,X2) => plus_0(succ_0(X0),succ_0(X1),X2))),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f80,plain,(
  ( ! [X0 : 'Nat_0()'] : (plus_0(X0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f8])).
tff(f8,axiom,(
  ! [X0 : 'Nat_0()'] : plus_0(X0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f83,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(zero_0,succ_0(X1),zero_0,X0)) )),
  inference(cnf_transformation,[],[f29])).
tff(f29,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : mulacc_0(zero_0,succ_0(X1),zero_0,X0)),
  inference(rectify,[],[f21])).
tff(f21,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : mulacc_0(zero_0,succ_0(X0),zero_0,X1)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f82,plain,(
  ( ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : (mulacc_0(zero_0,zero_0,X1,X0)) )),
  inference(cnf_transformation,[],[f28])).
tff(f28,plain,(
  ! [X0 : 'Nat_0()',X1 : 'Nat_0()'] : mulacc_0(zero_0,zero_0,X1,X0)),
  inference(rectify,[],[f22])).
tff(f22,axiom,(
  ! [X1 : 'Nat_0()',X0 : 'Nat_0()'] : mulacc_0(zero_0,zero_0,X0,X1)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f81,plain,(
  ( ! [X0 : 'Nat_0()'] : (addacc_0(X0,zero_0,zero_0,X0)) )),
  inference(cnf_transformation,[],[f11])).
tff(f11,axiom,(
  ! [X0 : 'Nat_0()'] : addacc_0(X0,zero_0,zero_0,X0)),
  file('/home/columpio/RiderProjects/RInGen/Tests/tests/generated/OriginalSolve/TIP.FreeSorts/tip2015/weird_nat_mul3acc_assoc3.smt2',unknown)).
tff(f8360,plain,(
  sP5(zero_0)),
  inference(unit_resulting_resolution,[],[f82,f8357,f112])).
tff(f8357,plain,(
  sP4(zero_0,zero_0)),
  inference(unit_resulting_resolution,[],[f211,f8335,f110])).
tff(f8335,plain,(
  sP3(zero_0,succ_0(succ_0(succ_0(zero_0))),zero_0)),
  inference(unit_resulting_resolution,[],[f81,f488,f108])).
tff(f211,plain,(
  ( ! [X0 : 'Nat_0()'] : (sP2(X0)) )),
  inference(unit_resulting_resolution,[],[f200,f106])).
